---
title: "Bivariate Description: Two Categorical (Nominal and/or Ordinal) Variables"
tutorial:
  id: "07-Bivariate-Description-Categorical-Data"
output: 
  QPATutorialsCourse::qpa_tutorial:
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
description: >
  Learn how to describe relationships involving categorical variables.
---

## Overview {#main-content}

Tutorials 4 through 6 described individual variables one at a time. Now we turn to the central task of empirical political science: [**bivariate description**]{.important-text} --- describing the relationship between two variables. Tutorials 7, 8, and 9 cover that together, each handling a different combination of variable types.

Tutorials 1 through 6 introduced tools one at a time and practiced them in isolation. Tutorials 7, 8, and 9 are organized differently. Each section starts with a research question, states a testable hypothesis, and then asks whether the pattern in the data is consistent with that hypothesis. This is how empirical political science works: you do not collect data and look for patterns at random --- you start with a question, derive a prediction from a theoretical explanation, and then evaluate whether the evidence points in the predicted direction. The tools you learned in Tutorials 1 through 6 are what you use to do that evaluation. In this tutorial and the next two, you will see how those tools fit into that larger workflow.

In this tutorial, you will learn how to describe the relationship between two variables measured at the nominal and/or ordinal level. Specifically, you will learn how to:

* Conduct a descriptive analysis of the relationship between two nominal or ordinal variables, or one of each, using cross-tabs
* Use the `CrossTable()` function in the [descr]{.package-name} package to produce cross-tabs
* Use column proportions to evaluate whether a pattern is consistent with a hypothesis
* Visualize cross-tab patterns with grouped bar plots using `ggplot()` with the `fill` aesthetic, `position = "dodge"`, and `after_stat(prop)`

```{r setup, include=FALSE}
library(learnr)
library(learnrhash)
library(tidyverse)
library(knitr)
library(poliscidata)
library(descr)
library(gradethis)
tutorial_options(exercise.checker = gradethis::grade_learnr,
                 exercise.completion = FALSE)
knitr::opts_chunk$set(error = TRUE)
data("counties", package = "QPATutorialsCourse")
counties$rural <- 1
counties$rural[counties$rural_urban <= 3] <- 0
counties$rural[counties$rural_urban >= 8] <- 2
data("world", package = "poliscidata")
```


## Bivariate Description

Once we have described the main features of the individual variables in our study, we want to look at the relationships between variables. The hypotheses we bring to our analyses suggest what we should expect to see. For example, if we hypothesize that more rural counties were more likely to support Donald Trump in 2016, we can examine the evidence by comparing the proportion of counties that cast a majority of their votes for Trump in more rural counties to that in more urban counties.

How we describe a bivariate relationship depends on the level of measurement of both variables.

- If both variables are categorical --- some combination of nominal and ordinal variables --- we describe the relationship using a cross-tabulation, which we may choose to visualize as a grouped bar plot.
- If one variable is categorical and one is interval, we compare the mean of the interval variable across categories of the categorical variable and may also compare the full distribution using histograms or box plots.
- If both variables are interval, we illustrate the relationship with a scatter plot and summarize its strength and direction with a correlation coefficient.

This tutorial covers the first case. Tutorials 8 and 9 cover the other two.

## Describing the Relationship Between Two Nominal or Ordinal Variables

We use a [**cross-tabulation**]{.important-text} --- or cross-tab --- to describe the relationship between two categorical variables. A cross-tab is a table that shows the distribution of cases across values of the outcome variable for cases that have different values on the explanatory variable. You've seen tables like this in textbooks and other contexts many times, even if you haven't called them cross-tabs before.

Each row of the table presents the different values the outcome variable can take, and each column represents the values the explanatory variable can take. Each cell gives the number of cases with that combination of row and column values. The margins give the row and column totals, and the bottom-right corner gives the total number of cases. A cross-tab may also report proportions calculated out of the row total, the column total, or all cases in the table. Which representation is most useful depends on what we want to compare --- and for evaluating hypotheses, we will almost always want column proportions.

The image below shows a cross-tab of the relationship between party identification and attitudes toward gay and lesbian relationships, with column proportions. Take a moment to orient yourself before reading the annotations --- find the row totals on the right, the column totals at the bottom, and try reading one cell on your own.

```{r crosstab-gay-support-by-party, echo=FALSE, out.width='70%', fig.alt='Annotated table showing a crosstabulation of support for gay marriage by party identification, with colored arrows explaining how to interpret specific cells. Rows list three positions: allowing gay and lesbian couples to legally marry, allowing civil unions but not marriage, and no legal recognition. Columns show respondents who identify as Democrat, Republican, Independent, and the overall total. The annotations highlight examples including the count of Democrats who support legal marriage, the percent of Independents who oppose legal recognition, the total number and percent supporting civil unions only, the total number of Democrats, and the overall sample size of respondents.'}
knitr::include_graphics("images/Crosstab2.jpg")
```

Reading the table: 1,196 respondents, or 51.8% *of Democratic respondents*, supported legal marriage; 385 respondents, or 21.8% *of Independents*, opposed any legal recognition; there were 2,309 Democratic respondents in total; and 5,492 respondents answered both questions. Notice that each percentage is calculated within its column --- that is what makes cross-tabs useful for comparing groups.

Cross-tabs are useful because they let us compare how common different outcomes are across groups. When the explanatory variable is in the columns and the outcome variable is in the rows, column proportions answer the question: within each category of the explanatory variable, what proportion of cases falls into each outcome category? That comparison helps us see whether the pattern in the data is consistent with what a hypothesis predicts.

### The `CrossTable()` Function

To produce a cross-tab in R, we use the `CrossTable()` function from the [descr]{.package-name} package. Because cross-tabs can be organized in more than one way, we follow the same three rules throughout this tutorial.

::: {.important-note}
**Rules for organizing cross-tabs**

Rule 1: The outcome variable goes in the rows of the table. In `CrossTable()`, this is the `x` argument.

Rule 2: The explanatory variable goes in the columns of the table. In `CrossTable()`, this is the `y` argument.

Rule 3: We use column proportions. Column proportions show how the outcome is distributed within each category of the explanatory variable.
:::

The `x` argument defines the variable that appears in the rows, so we use it for the outcome or dependent variable. The `y` argument defines the variable that appears in the columns, so we use it for the explanatory or independent variable. By default, `CrossTable()` reports row proportions, column proportions, and total proportions all at once, which makes the table harder to read. Set `prop.c = TRUE` to keep column proportions and `prop.r = FALSE`, `prop.t = FALSE`, and `prop.chisq = FALSE` to suppress everything else.


## Are Rural Counties More Likely to Support Trump?

Let’s start with a familiar political question. Rural and urban counties often differ in their economic conditions, social composition, and political priorities. If Donald Trump's campaign was more successful in rural areas, then counties that are more rural should be more likely to have given Trump a majority of their votes.

::: {.hypothesis}
In a comparison of [counties], those that are [more rural] will be more likely to [give a majority of their support to Trump] than will those that are [less rural (more urban)].
:::

The data frame **counties** contains a variable called **rural**, a 3-category ordinal variable where 0 represents more urban counties, 1 represents counties that are neither urban nor rural, and 2 represents more rural counties. It also contains a nominal variable called **TrumpMajority**, which takes a value of 0 if Hillary Clinton won a majority of the vote and 1 if Donald Trump won a majority of the vote. We will use these variables to create our first cross-tab.

### Create the Cross-tab

We need to first load the [descr]{.package-name} package with the `library()` function so that we can use the `CrossTable()` function. For this cross-tab, **TrumpMajority** is the outcome variable, so it goes in the `x` argument and will appear in the rows. **rural** is the explanatory variable, so it goes in the `y` argument and will appear in the columns. We set `prop.c = TRUE` because we want column proportions: the proportion of counties within each rurality category that gave a majority of votes to Clinton or Trump.

**Run and observe: Generate a cross-tab with column proportions**

```{r crosstab-trump-majority-by-rural-column-proportions, exercise=TRUE, exercise.cap="Run and observe: Cross-tab of TrumpMajority by rural status"}
library(descr) 
CrossTable(x = counties$TrumpMajority,
           y = counties$rural,
           prop.r = FALSE,
           prop.c = TRUE,
           prop.t = FALSE,
           prop.chisq = FALSE)
```

Once a package is loaded with `library()`, all of its functions are available for the rest of your R session --- you do not need to load it again each time you use a function from it. Since [descr]{.package-name} was loaded in the exercise above, you can assume it is already loaded for all subsequent exercises in this tutorial.

**Practice: Read the cross-tab**

Use the column totals and cell proportions in the table to answer the next two questions.


```{r quiz-identify-number-of-counties, echo=FALSE}
question("How many counties in the sample are rural?",
  correct = "Right. 627 counties in the sample are rural.",
  answer("2622", message = "2622 is a row total: it is the number of counties where Trump won a majority, not the number of rural counties."),
  answer("1162", message = "1162 is a column total, but for the wrong column. Urban counties are coded 0, and 1162 is the total for that column."),
  answer("627", correct = TRUE, message = "627 is the column total for rural counties --- those coded 2 on the rural variable. Column totals at the bottom of the table give the number of counties in each category of the explanatory variable."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  try_again = "Incorrect. Try again."
)
```


```{r quiz-identify-county-trump-majority, echo=FALSE}
question("What does the value 0.261 in the top-left cell of the table tell us?",
  correct = "Right. That is a column proportion within the Urban column.",
  answer("The proportion of all counties that cast votes for Hillary Clinton",
         message = "This would be a total proportion calculated out of all counties. Column proportions are calculated within each column, not across all counties."),
  answer("The proportion of urban counties that cast a majority of the vote for Donald Trump",
         message = "The top row corresponds to counties where Clinton won, not Trump. Look at which row this cell is in."),
  answer("The proportion of urban counties that cast a majority of the vote for Hillary Clinton",
         correct = TRUE,
         message = "The top-left cell is in the Clinton-majority row (TrumpMajority = 0) and the Urban column (rural = 0). The column proportion 0.261 means 26.1% of urban counties gave Clinton a majority."),
  answer("The proportion of counties that cast a majority of votes for Hillary Clinton that were also urban counties",
         message = "That would be a row proportion --- calculated out of all Clinton-majority counties. Column proportions are calculated within each column, not within each row."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  try_again = "Incorrect. Try again."
)
```


### Is the Pattern Consistent with the Hypothesis?

Reading individual cells is the first step. The bigger goal is to use the table to see whether the overall pattern matches what the hypothesis predicted. Because **TrumpMajority** is the outcome variable, we focus on the row where **TrumpMajority** = 1. Because **rural** is the explanatory variable, we compare the column proportions across the urban, neither urban nor rural, and rural columns.

The hypothesis predicts that the proportion of counties where Trump won a majority should increase as counties become more rural. We would also expect the reverse pattern in the Clinton-majority row: the proportion of Clinton-majority counties should decrease as counties become more rural.

::: {.caution}
**Cross-tab analysis is not a hypothesis test.** A cross-tab can show whether the observed pattern is *consistent with* a hypothesis. It does not tell us whether the difference is larger than we would expect by chance. Formal hypothesis tests are covered in later tutorials.
:::

The cross-tab is reproduced here so you do not have to look back.

```{r crosstab-trump-majority-by-rural-column-proportions-repeat, exercise=FALSE, echo=FALSE}
CrossTable(x = counties$TrumpMajority,
           y = counties$rural, 
           prop.r = FALSE, 
           prop.c = TRUE, 
           prop.t = FALSE,
           prop.chisq = FALSE)
```


**Practice: Connect the pattern to the hypothesis**

Use the Trump-majority row to decide whether the observed pattern matches what the hypothesis predicted.

```{r quiz-pattern-hypothesis, echo=FALSE}
question("What happens to the proportion of Trump-majority counties as counties become more rural?",
  correct = "Right. The proportion of Trump-majority counties increases with rurality.",
  answer("The proportion increases as counties become more rural",
         correct = TRUE,
         message = "Reading across the Trump-majority row, the proportion rises from urban (0.739) to neither urban nor rural (0.891) to rural (0.931) --- exactly the monotonic pattern the hypothesis predicted."),
  answer("The proportion decreases as counties become more rural",
         message = "Look across the Trump-majority row from left to right. The proportions get larger, not smaller."),
  answer("The proportion is about the same across all three rurality categories",
         message = "The proportions are not about the same. The difference between urban counties and rural counties is nearly 20 percentage points."),
  answer("The cross-tab only shows counts, so we cannot compare proportions",
         message = "This table includes column proportions. Those proportions are the entries we use to compare across rurality categories."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  try_again = "Incorrect. Try again."
)
```


```{r quiz-rural-hypothesis-eval, echo = FALSE}
question("Based on that pattern, which statement best describes whether the cross-tab is consistent with the hypothesis that more rural counties were more likely to support Trump?",
  correct = "Right. The pattern is consistent with the hypothesis.",
  answer("The pattern is consistent with the hypothesis --- the proportion of counties giving Trump a majority increases from 0.739 to 0.891 to 0.931 as counties become more rural",
         correct = TRUE,
         message = "The hypothesis predicts that Trump-majority counties should become more common as rurality increases, and that is exactly what the column proportions show."),
  answer("The pattern is not consistent with the hypothesis because Trump won a majority of counties in every rurality category",
         message = "The hypothesis is not about whether Trump won more than half of counties in every category. It is about whether the proportion of Trump-majority counties increases as counties become more rural --- and it does, from 0.739 to 0.931."),
  answer("The pattern is not consistent with the hypothesis because the differences across columns are too small to describe",
         message = "The difference between urban counties (0.739) and rural counties (0.931) is nearly 20 percentage points, which is a meaningful descriptive difference."),
  answer("The cross-tab cannot tell us whether the pattern is consistent with the hypothesis without knowing the total number of counties",
         message = "Column proportions allow us to compare categories even when the column totals differ. The counts are useful context, but the proportions are the key entries for evaluating this descriptive pattern."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  try_again = "Incorrect. Try again."
)
```


## Are Majority-White Counties More Likely to Support Trump?

Racial composition may be associated with local political context, party coalitions, and the social environments in which campaigns are received. If Donald Trump's campaign was more successful in majority-white counties, then majority-white counties should be more likely than other counties to have given Trump a majority of their votes.

::: {.hypothesis}
In a comparison of [counties], those that are [majority-white] will be more likely to [give a majority of their support to Trump] than will those that are [not majority-white].
:::

The data frame **counties** contains a nominal variable called **racial_majority** that takes the character values "Black Majority," "Hispanic Majority," "Other/No Majority," and "White Majority." We will use **racial_majority** as the explanatory variable and **TrumpMajority** as the outcome variable.

### Create the Cross-tab

To evaluate whether the observed pattern is consistent with this hypothesis, we need to compare Trump support across racial majority categories --- seeing whether the proportion of Trump-majority counties is substantially higher for White Majority counties than for the others. A cross-tab with column proportions gives us exactly that comparison.

**Practice: Create a cross-tab of Trump support by racial majority**

Apply the same three rules: outcome in rows, explanatory in columns, column proportions only. Two hints are available. Click "Submit Answer" to check your work.

```{r exercise-create-crosstab-trump-racial, exercise=TRUE, exercise.cap="Practice: Cross-tab of TrumpMajority by racial_majority"}
CrossTable(x = counties$___,
           y = counties$___,
           prop.r = ___,
           prop.c = ___,
           prop.t = ___,
           prop.chisq = ___)
```

```{r exercise-create-crosstab-trump-racial-hint-1}
# Hint 1 of 2
# Apply the three rules: the outcome goes in x (rows),
# the explanatory variable goes in y (columns).
# Which variable is the outcome and which is explanatory here?
```

```{r exercise-create-crosstab-trump-racial-hint-2}
# Hint 2 of 2
# x = counties$TrumpMajority (the outcome, in rows)
# y = counties$racial_majority (the explanatory variable, in columns)
# You want column proportions only --- set prop.c = TRUE
# and all other prop arguments to FALSE to keep the table clean.
```

```{r exercise-create-crosstab-trump-racial-check}
grade_this({
  code <- paste(.user_code, collapse = "\n")
  code_no_comments <- gsub("#.*", "", code)

  if (!grepl("CrossTable\\s*\\(", code_no_comments)) {
    fail("Use the CrossTable() function to create the cross-tab.")
  }
  if (!grepl("x\\s*=\\s*counties\\$TrumpMajority", code_no_comments)) {
    fail("TrumpMajority should be the x argument --- it is the outcome variable and belongs in the rows.")
  }
  if (!grepl("y\\s*=\\s*counties\\$racial_majority", code_no_comments)) {
    fail("racial_majority should be the y argument --- it is the explanatory variable and belongs in the columns.")
  }
  if (!grepl("prop\\.c\\s*=\\s*TRUE", code_no_comments)) {
    fail("Set prop.c = TRUE to show column proportions.")
  }
  if (!grepl("prop\\.r\\s*=\\s*FALSE", code_no_comments)) {
    fail("Set prop.r = FALSE to suppress row proportions.")
  }
  if (!grepl("prop\\.t\\s*=\\s*FALSE", code_no_comments)) {
    fail("Set prop.t = FALSE to suppress total proportions.")
  }
  if (!grepl("prop\\.chisq\\s*=\\s*FALSE", code_no_comments)) {
    fail("Set prop.chisq = FALSE to suppress chi-square contributions and reduce clutter.")
  }

  pass("Your cross-tab is correct --- TrumpMajority is in the rows because it is the outcome variable in our hypothesis, and racial_majority is in the columns because it is the explanatory variable. This orientation lets us compare column proportions across racial majority categories to see how the outcome varies. Examine those proportions before answering the next question.")
})
```

Two counties have missing values on **racial_majority**, so the cross-tab is based on the 3,110 counties with racial-majority data.

### Interpreting the Cross-tab

Use the column totals and cell proportions in the table to answer the next question.

```{r quiz-identify-crosstab-cells, echo=FALSE}
question("Which of the following statements about the cross-tab are true? Select all that apply.",
  correct = "Right. Both statements correctly describe column proportions.",
  answer("Of the 490 counties that cast a majority of votes for Hillary Clinton, 0.111 were majority white",
         message = "This reverses the direction. The proportion 0.111 is a column proportion calculated out of the 2866 White Majority counties --- not out of Clinton-majority counties."),
  answer("Of the 2866 counties that were majority white, 0.111 or 11.1% of them cast a majority of their votes for Hillary Clinton",
         correct = TRUE,
         message = "The White Majority column proportion for Clinton-majority counties is 0.111. That means 11.1% of majority-white counties gave Clinton a majority."),
  answer("Of the 3,110 counties with racial-majority data, 0.039 or 3.9% cast a majority of their votes for Donald Trump",
         message = "This would be a total proportion calculated out of all counties with racial-majority data. The value 0.039 is a column proportion within the Black Majority column."),
  answer("Of the 102 majority Black counties, 0.039 or 3.9% cast a majority of their vote for Donald Trump",
         correct = TRUE,
         message = "The Black Majority column proportion for Trump-majority counties is 0.039. That means only 3.9% of majority-Black counties gave Trump a majority."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  try_again = "Incorrect. Try again."
)
```


### Is the Pattern Consistent with the Hypothesis?

To ask whether the pattern is consistent with the hypothesis, focus on the Trump-majority row and compare the column proportions across the racial-majority categories. The hypothesis predicts a particular pattern across those columns.

Remember, this is a descriptive comparison. The cross-tab can show whether the observed pattern is consistent with the hypothesis, but it does not explain why the pattern exists or tell us whether the relationship is statistically significant.

**Practice: Connect the pattern to the hypothesis**

Use the Trump-majority row to answer the next question.

```{r quiz-racial-hypothesis, echo=FALSE}
question("Which statement best describes whether the cross-tab is consistent with the hypothesis that majority-white counties were more likely to support Trump?",
  correct = "Right. The pattern is consistent with the hypothesis.",
  answer("The pattern is consistent with the hypothesis --- White Majority counties have the highest proportion supporting Trump (0.889) of any racial majority category",
         correct = TRUE,
         message = "The White Majority column shows 0.889 of counties giving Trump a majority --- higher than any other racial majority category. This is consistent with the hypothesis."),
  answer("The pattern is not consistent with the hypothesis --- Trump won a majority in both White Majority and Hispanic Majority counties",
         message = "Look at the proportions carefully. The proportion of Hispanic Majority counties giving Trump a majority is much lower than for White Majority counties. The hypothesis is about relative proportions, not just whether Trump won more than 50%."),
  answer("The pattern is not consistent with the hypothesis because the table does not explain why counties differ",
         message = "A cross-tab does not explain why the pattern exists. Here we are only asking whether the descriptive pattern matches what the hypothesis predicted."),
  answer("The cross-tab is uninformative because the columns have very different numbers of counties",
         message = "That is exactly why we use column proportions. By dividing by the column total, we standardize the comparison so that different column sizes do not distort the results."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  try_again = "Incorrect. Try again."
)
```

## Are Dictatorships Less Likely to Have High Gender Equality?

Democratic institutions may create more pressure for equal rights, representation, and policy responsiveness. Dictatorships may face fewer pressures to protect gender equality. For that reason, we might expect dictatorships to be less likely than democratic regimes to have high levels of gender equality.

::: {.hypothesis}
In a comparison of [countries], those that are [dictatorships] will be less likely to [have a high degree of gender equality] than will those that are [not dictatorships].
:::

The data frame **world** contains the nominal variable **regime_type3**, which takes the values "Dictatorship," "Parliamentary democ," and "Presidential democ." It also contains the ordinal variable **gender_equal3**, which may be "Low," "Medium," or "High." We will use **regime_type3** as the explanatory variable and **gender_equal3** as the outcome variable.

### Create the Cross-tab

To evaluate whether the observed pattern is consistent with this hypothesis, we need to compare the distribution of gender equality across regime types --- seeing whether the proportion of countries with high gender equality is lower for dictatorships than for democratic regimes. A cross-tab with column proportions gives us that comparison directly.

**Practice: Create a cross-tab of gender equality by regime type**

Write the code to produce a cross-tab of **gender_equal3** by **regime_type3** from the **world** data frame, following the same three rules as before. Three hints are available. Click "Submit Answer" to check your work.

```{r exercise-create-crosstab-gender-regime, exercise=TRUE, exercise.cap="Practice: Cross-tab of gender equality by regime type"}

```

```{r exercise-create-crosstab-gender-regime-hint-1}
# Hint 1 of 3
# Use the CrossTable() function with x as the outcome
# variable and y as the explanatory variable.
# What is the outcome here, and what is the explanatory variable?
```

```{r exercise-create-crosstab-gender-regime-hint-2}
# Hint 2 of 3
# The outcome is gender equality (world$gender_equal3).
# The explanatory variable is regime type (world$regime_type3).
# Set prop.c = TRUE and all other prop arguments to FALSE.
```

```{r exercise-create-crosstab-gender-regime-hint-3}
# Hint 3 of 3
# CrossTable(x = world$gender_equal3,
#            y = world$regime_type3, ...)
# gender_equal3 is the outcome --- it goes in x (rows).
# regime_type3 is the explanatory variable --- it goes in y (columns).
# Set prop.c = TRUE and all other prop arguments to FALSE.
```

```{r exercise-create-crosstab-gender-regime-check}
grade_this({
  code <- paste(.user_code, collapse = "\n")
  code_no_comments <- gsub("#.*", "", code)

  if (!grepl("CrossTable\\s*\\(", code_no_comments)) {
    fail("Use the CrossTable() function to create the cross-tab.")
  }
  if (!grepl("x\\s*=\\s*world\\$gender_equal3", code_no_comments)) {
    fail("gender_equal3 should be the x argument --- it is the outcome variable and belongs in the rows.")
  }
  if (!grepl("y\\s*=\\s*world\\$regime_type3", code_no_comments)) {
    fail("regime_type3 should be the y argument --- it is the explanatory variable and belongs in the columns.")
  }
  if (!grepl("prop\\.c\\s*=\\s*TRUE", code_no_comments)) {
    fail("Set prop.c = TRUE to show column proportions.")
  }
  if (!grepl("prop\\.r\\s*=\\s*FALSE", code_no_comments)) {
    fail("Set prop.r = FALSE to suppress row proportions.")
  }
  if (!grepl("prop\\.t\\s*=\\s*FALSE", code_no_comments)) {
    fail("Set prop.t = FALSE to suppress total proportions.")
  }
  if (!grepl("prop\\.chisq\\s*=\\s*FALSE", code_no_comments)) {
    fail("Set prop.chisq = FALSE to suppress chi-square contributions and reduce clutter.")
  }

  pass("Your cross-tab is correct --- gender_equal3 is in the rows because it is the outcome variable, and regime_type3 is in the columns because it is the explanatory variable. This lets us compare column proportions across regime types to see how gender equality varies. Examine those proportions before answering the next question.")
})
```

### Interpreting the Cross-tab

**Practice: Read the cross-tab**

Use the column proportions in the table to answer the next question.


```{r quiz-identify-proportions-gender-regime, echo=FALSE}
question("Which of the following statements are true? Select all that apply.",
  correct = "Right. Both statements correctly describe column proportions.",
  answer("Of the countries with a regime type of Dictatorship, the level of gender equality with the highest proportion of cases is low (0.625)",
         correct = TRUE,
         message = "The Dictatorship column shows that 0.625 of dictatorships have low gender equality --- the highest proportion in that column."),
  answer("Of the 21 countries with the highest level of gender equality, only 0.062 are dictatorships",
         message = "This reverses the direction. The value 0.062 is a column proportion calculated out of all dictatorships --- not out of the 21 high-equality countries."),
  answer("Parliamentary democracies have the highest proportion of countries with high gender equality",
         correct = TRUE,
         message = "Looking across the High row, the Parliamentary democracy column has the highest proportion. Parliamentary democracies are the most likely regime type to have high gender equality."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  try_again = "Incorrect. Try again."
)
```


### Is the Pattern Consistent with the Hypothesis?

To ask whether the pattern is consistent with the hypothesis, focus on the High gender equality row and compare the column proportions across regime types. The hypothesis predicts that dictatorships should be less likely than democratic regimes to have high gender equality.

**Practice: Connect the pattern to the hypothesis**

Use the High gender equality row to decide whether the observed pattern matches what the hypothesis predicted.

```{r quiz-gender-hypothesis, echo=FALSE}
question("Which statement best describes whether the cross-tab is consistent with the hypothesis that dictatorships are less likely to have high gender equality?",
  correct = "Right. The pattern is consistent with the hypothesis.",
  answer("The pattern is consistent with the hypothesis --- dictatorships have a lower proportion with high gender equality than both democratic regime types",
         correct = TRUE,
         message = "Only 0.062 of dictatorships have high gender equality, compared to 0.200 of presidential democracies and 0.448 of parliamentary democracies. Dictatorships are less likely to have high gender equality than either type of democracy, consistent with the hypothesis."),
  answer("The pattern is not consistent with the hypothesis --- some dictatorships do have high gender equality",
         message = "The hypothesis does not require that zero dictatorships have high gender equality. It predicts that dictatorships are less likely to have high gender equality than non-dictatorships."),
  answer("The cross-tab is uninformative because the columns have different numbers of countries",
         message = "That is exactly why we use column proportions. By dividing by the column total, we standardize the comparison so that different column sizes do not distort the results."),
  answer("The pattern is partially consistent with the hypothesis --- it holds for parliamentary democracies but not for presidential democracies",
         message = "Look at the High gender equality row across all three columns. Both democratic regime types have higher proportions of high gender equality than dictatorships, which is consistent with the hypothesis."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  try_again = "Incorrect. Try again."
)
```


## Visualizing Cross-tabs with Grouped Bar Plots

Cross-tabs give us the precise numbers. Grouped bar plots help us see the same pattern visually. A grouped bar plot uses one variable on the x-axis and uses fill color to show categories of the second variable. The height of the bars can show either counts or proportions.

The plot version follows the same organization we used in the cross-tab. The variable that was in the rows of the cross-tab goes on the x-axis, and the variable that was in the columns of the cross-tab becomes the `fill` variable. For the Trump/rural example, **TrumpMajority** goes on the x-axis and **rural** is represented by fill color. This may feel different from some bar plots you have seen, where the explanatory variable is placed on the x-axis. Here we keep the outcome on the x-axis because we want the plot to mirror the cross-tab logic: compare outcome proportions across categories of the explanatory variable.

### The `fill` Aesthetic and Side-by-Side Bars

Building a grouped bar plot uses the same basic structure as the single-variable bar plots from Tutorials 4 and 5, with two additions.

**`fill = factor(rural)`** inside the main `aes()` call tells ggplot to color the bars by the categories of **rural**. This is a different use of `fill` than you have seen before --- here it maps a variable to a visual property rather than setting a single fixed color. The `factor()` wrapper ensures ggplot treats the numeric codes as discrete categories rather than points on a continuous scale. Both `x` and `fill` are wrapped in `factor()` here because both **TrumpMajority** and **rural** are stored as numbers. When a variable is already stored as a character (like **racial_majority**), `factor()` is not needed.

**`position = "dodge"`** inside `geom_bar()` places the fill categories side by side within each x-axis group rather than stacking them on top of each other. Stacked bars show how the total is composed; side-by-side bars make it easy to compare heights across fill categories, which is what we need to evaluate a hypothesis.

**Run and observe: Create a side-by-side bar plot with position = "dodge"**

```{r barplot-dodge-trump-majority-rural, exercise=TRUE, exercise.cap="Run and observe: Side-by-side bar plot with position = \"dodge\""}
library(ggplot2)
counties %>%
  filter(!is.na(TrumpMajority), !is.na(rural)) %>%
  ggplot(mapping = aes(x = factor(TrumpMajority), fill = factor(rural))) +
  geom_bar(position = "dodge") 
```

Since [ggplot2]{.package-name} was loaded above, you can assume it is already loaded for all subsequent exercises in this tutorial.

This is a count plot. The x-axis shows the outcome categories (Clinton majority = 0, Trump majority = 1), and the fill colors show the three rurality categories placed side by side within each outcome group. The bars show raw counts, so their heights depend partly on how many counties fall in each rurality category. Because there are very different numbers of urban, neither-urban-nor-rural, and rural counties, counts alone can be misleading for comparing voting patterns. Proportions give a clearer picture.

As established in earlier tutorials, always filter missing values in the variables used in a plot before plotting. This keeps missing categories out of the display and avoids ggplot dropping observations or producing warnings. Here, **TrumpMajority** and **rural** have no missing values, so the filter changes nothing --- but the habit applies consistently.

### Plotting Proportions

The count plot is a useful starting point, but counts are misleading when the groups being compared have different sizes --- and here, urban, neither-urban-nor-rural, and rural counties have very different counts. Proportions put all three groups on the same 0-to-1 scale, giving the visual equivalent of column proportions in the cross-tab.

Two new arguments inside `geom_bar()` make this work:

- **`y = after_stat(prop)`** inside `aes()` transforms the y-axis from raw counts to proportions. The `after_stat()` function computes a new quantity from the data after the bars have been tallied. Setting it to `prop` asks ggplot to divide each bar's count by the group total rather than showing the raw count.
- **`group = factor(rural)`** tells ggplot which variable to use when calculating proportions. Without it, ggplot would calculate proportions out of the entire dataset, so each bar would show what fraction of all counties falls in that outcome category --- not what fraction within each rurality group. Setting `group = factor(rural)` makes the Clinton and Trump bars within each rurality group add to 1, giving the visual equivalent of the column proportions in the cross-tab.

Both arguments go inside a second `aes()` call within `geom_bar()`, not in the main `ggplot()` call. They belong here because they describe how the bars are computed, not how the overall data is mapped.

**Run and observe: Plot proportions instead of frequencies**

```{r barplot-proportions-trump-majority-rural, exercise=TRUE, exercise.cap="Run and observe: Proportion bar plot with after_stat(prop)"}
counties %>%
  filter(!is.na(TrumpMajority), !is.na(rural)) %>%
  ggplot(mapping = aes(x = factor(TrumpMajority), fill = factor(rural))) +
  geom_bar(position = "dodge", mapping = aes(y = after_stat(prop), group = factor(rural))) 
```

Now the y-axis shows proportions rather than counts. For each rurality category, the Clinton and Trump proportions add to 1. This lets us compare the share of counties in each rurality category that gave a majority of votes to Clinton or Trump. In the plot, we can see that the proportion of Trump-majority counties is highest among rural counties and lowest among urban counties.

**Practice: Why use proportions instead of frequencies?**

Think about what happens when you compare counts across groups of very different sizes before answering.

```{r quiz-proportions-vs-frequencies, echo=FALSE}
question("Why is the proportion bar plot more useful than the frequency bar plot for comparing Clinton and Trump support across urbanicity categories?",
  correct = "Right. Proportions remove the group-size effect that makes raw counts misleading.",
  answer("Because proportions are always easier to read than counts",
         message = "Proportions are not always preferable. When you want to communicate how many cases are in a category, counts are more direct. The issue here is specific to comparison across groups of different sizes."),
  answer("Because the three urbanicity categories have different numbers of counties, so raw counts make the comparison misleading", correct = TRUE,
         message = "There are 1162 urban counties, 1323 neither counties, and 627 rural counties. Comparing raw counts across these groups would mix up differences in group size with differences in voting patterns. Proportions remove the group-size effect."),
  answer("Because ggplot cannot display frequencies for grouped bar plots",
         message = "ggplot can display frequencies for grouped bar plots --- that is what the first two Try it exercises produced. The choice between frequencies and proportions is a substantive one, not a technical limitation."),
  answer("Because proportions must sum to one, which makes the math easier to check",
         message = "While proportions within each fill category do sum to one here, that is a consequence of the approach rather than the reason to use it. The real reason is about comparability across groups of different sizes."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  try_again = "Incorrect. Try again."
)
```

### Making the Plot More Informative

The plot above shows the right information, but it is not yet easy to read. The x-axis uses the values 0 and 1 instead of candidate names, and the legend uses the values 0, 1, and 2 instead of meaningful labels for rurality. We can make the plot more informative by adding labels.

We will use `labs()` to add a title, axis labels, and a caption. We will use `scale_x_discrete()` to change the x-axis labels from 0 and 1 to "Clinton" and "Trump." We will also use `scale_fill_discrete()` to make the legend more informative. The `name` argument gives the legend a title, and the `labels` argument gives readable names to the values of the `fill` variable.

Before you execute the code below, read through it carefully. You will have to edit similar code in later exercises.

**Run and observe: Add labels to make the plot easier to read**

```{r barplot-informative-trump-majority-rural, exercise=TRUE, exercise.cap="Run and observe: Adding labels with labs() and scale_fill_discrete()"}
counties %>%
  filter(!is.na(TrumpMajority), !is.na(rural)) %>%
  ggplot(mapping = aes(x = factor(TrumpMajority), fill = factor(rural))) +
  geom_bar(position = "dodge", aes(y = after_stat(prop), group = factor(rural))) +
  labs(title = "Proportion of Counties Casting \na Majority of Votes for Clinton and Trump", 
       x = NULL, 
       y = "Proportion",
       caption = "Source: Linn, Nagler, Zilinsky") +
  scale_x_discrete(labels = c("0" = "Clinton", "1" = "Trump")) +
  scale_fill_discrete(name = "Urbanicity", 
                       labels =  c("Urban", "Neither Urban \nnor Rural", "Rural")) 
```


### Plot Vote Outcome by Racial Majority

Now we can use the same plotting structure with the racial-majority cross-tab. We want to present the proportion of counties casting a majority of votes for Hillary Clinton and Donald Trump grouped by the categories of **racial_majority**. The outcome variable, **TrumpMajority**, goes on the x-axis. The explanatory variable, **racial_majority**, is represented by fill color. Because **TrumpMajority** is stored as a numeric variable, we use `factor(TrumpMajority)` so that `ggplot()` treats it as categorical.

Two counties have missing racial-majority data. Following the norm established in earlier tutorials, filter before plotting so missing racial-majority values do not appear as an additional category alongside the substantive groups.

**Practice: Create a proportion bar plot of vote outcome by racial majority**

Use the same outcome/explanatory-variable logic from the cross-tab to decide which variable goes on the x-axis and which goes in `fill`, filter out the two counties with missing racial-majority data, then build the proportion bar plot. Four hints are available. Click "Submit Answer" to check your work.

```{r exercise-create-proportions-vote-race-majority, exercise=TRUE, exercise.cap="Practice: Proportion bar plot of vote outcome by racial majority"}
counties %>%
  filter(___) %>%
  ggplot(mapping = ___(x = factor(___), fill = ___)) +
  geom_bar(___ = ___, aes(y = after_stat(prop), group = ___)) +
  labs(title = "Proportion of Counties Casting \na Majority of Votes for Clinton and Trump", 
       x = NULL, 
       y = "Proportion",
       caption = "Source: Linn, Nagler, Zilinsky") +
  scale_x_discrete(labels = c("Clinton", "Trump")) +
  scale_fill_discrete(name = "Majority",
                      labels = c("Black", "Hispanic", "Other/No", "White")) 
```

```{r exercise-create-proportions-vote-race-majority-hint-1}
# Hint 1 of 4
# Fill in the filter() blank with !is.na(racial_majority)
# to remove the two counties with missing racial majority data.
# Then pipe into ggplot() and map the variables.
```

```{r exercise-create-proportions-vote-race-majority-hint-2}
# Hint 2 of 4
# Inside aes(), x is the outcome variable and
# fill is the grouping (explanatory) variable.
# x = factor(TrumpMajority) --- wrapped in factor()
# because TrumpMajority is stored as numeric.
# fill = racial_majority --- no factor() needed
# because it is already a character variable.
```

```{r exercise-create-proportions-vote-race-majority-hint-3}
# Hint 3 of 4
# Inside geom_bar(), set position = "dodge"
# to place bars side-by-side.
# The group argument inside aes() should match
# the fill variable: group = racial_majority.
```

```{r exercise-create-proportions-vote-race-majority-hint-4}
# Hint 4 of 4
# The full structure after filtering is:
# ggplot(mapping = aes(x = factor(TrumpMajority), fill = racial_majority)) +
# geom_bar(position = "dodge", aes(y = after_stat(prop), group = racial_majority))
# Then add labs() and scale_fill_discrete() for labels.
```

```{r exercise-create-proportions-vote-race-majority-check}
grade_this({
  code <- paste(.user_code, collapse = "\n")
  code_no_comments <- gsub("#.*", "", code)

  if (!inherits(.result, "ggplot")) {
    fail("Your code should return a ggplot object. Check for a missing + at the end of a line, unmatched parentheses, a label string missing its quotation marks, or a missing %>% in the pipeline.")
  }
  if (!grepl("filter\\s*\\(\\s*!\\s*is\\.na\\s*\\(\\s*racial_majority\\s*\\)", code_no_comments)) {
    fail("Filter out missing values with filter(!is.na(racial_majority)) before plotting.")
  }
  if (!grepl("x\\s*=\\s*factor\\s*\\(\\s*TrumpMajority\\s*\\)", code_no_comments)) {
    fail("Map factor(TrumpMajority) to the x-axis inside aes().")
  }
  if (!grepl("fill\\s*=\\s*racial_majority", code_no_comments)) {
    fail("Map racial_majority to the fill aesthetic inside aes().")
  }
  if (!grepl("geom_bar\\s*\\(", code_no_comments)) {
    fail("Use geom_bar() to create the bar plot.")
  }
  if (!grepl("position\\s*=\\s*['\"]dodge['\"]", code_no_comments)) {
    fail('Add position = "dodge" inside geom_bar() to place bars side-by-side.')
  }
  if (!grepl("y\\s*=\\s*after_stat\\s*\\(\\s*prop\\s*\\)", code_no_comments)) {
    fail("Set y = after_stat(prop) inside the aes() call within geom_bar() to plot proportions.")
  }
  if (!grepl("group\\s*=\\s*racial_majority", code_no_comments)) {
    fail("Set group = racial_majority inside the aes() call within geom_bar() so proportions are calculated within each racial majority category.")
  }

  # Check plot labels from object
  plot_labels <- .result$labels
  code_all <- paste(.user_code, collapse = "\n")
  if (!grepl("Clinton", code_all) || !grepl("Trump", code_all)) {
    fail('Add scale_x_discrete() with both "Clinton" and "Trump" as labels to replace the 0 and 1 x-axis values.')
  }
  if (!grepl("Black", code_all) || !grepl("Hispanic", code_all) ||
      !grepl("Other", code_all) || !grepl("White", code_all)) {
    fail('Add scale_fill_discrete() with all four racial majority category labels: "Black", "Hispanic", "Other/No", and "White".')
  }
  if (!identical(plot_labels$caption, "Source: Linn, Nagler, Zilinsky")) {
    fail('Set caption = "Source: Linn, Nagler, Zilinsky" inside labs().')
  }

  pass("Excellent! The bar plot makes the pattern from the cross-tab easier to see. Compare the Trump bar across racial-majority categories: White Majority counties have the highest Trump-majority proportion, while Black Majority counties have the lowest.")
})
```

This visualization presents the same information as the racial-majority cross-tab, but the pattern is easier to see. The Trump-majority proportion is highest for White Majority counties and lowest for Black Majority counties, consistent with the pattern you saw in the table.


### Plot Gender Equality by Regime Type

Now we can visualize the gender-equality/regime-type cross-tab. Plotting the proportions side by side lets us see at a glance whether the distribution of gender equality across regime types matches the pattern the hypothesis predicted. Both **gender_equal3** and **regime_type3** have missing values, so filter both before plotting.

Before writing the labels, check the order in which the categories are stored --- replacement labels in `scale_fill_discrete()` must be listed in the same order as the stored values.

**Run and observe: Check the stored order of regime_type3**

```{r regime-type-levels, exercise=TRUE, exercise.cap="Run and observe: Check category order for regime_type3"}
levels(world$regime_type3)
```

The output shows the order you must follow when listing replacement labels in `scale_fill_discrete()`.

::: {.tip}
**Line breaks in titles and labels:** When a label is too long to fit on one line, you can insert a line break using `\n` inside the quoted string. For example, `"Parliamentary\nDemocracy"` would display "Parliamentary" on one line and "Democracy" on the next. This works anywhere you use quoted text in `labs()` or `scale_fill_discrete()`.
:::

**Practice: Build the complete proportion bar plot for gender equality by regime type**

The variables are **gender_equal3** (outcome), **regime_type3** (explanatory), and the data frame is **world**. Use the same outcome/explanatory-variable logic to decide which goes on `x` and which goes in `fill`. The `filter()` step is provided --- complete the rest. Your plot should have:

- Title: "Gender Equality"
- y-axis label: "Proportion"
- x-axis label: omitted (`NULL`)
- Caption: "Country Data from the United Nations"
- Legend title: "Regime Type"
- Category labels: `"Dictatorship"`, `"Parliamentary\nDemocracy"`, `"Presidential\nDemocracy"` --- the `\n` in the two democracy labels creates a line break to keep the legend readable

Four hints are available. Click "Submit Answer" to check your work.

```{r exercise-create-proportions-gender-equality-regime, exercise=TRUE, exercise.cap="Practice: Proportion bar plot of gender equality by regime type"}
world %>%
  filter(!is.na(gender_equal3), !is.na(regime_type3)) %>%
```

```{r exercise-create-proportions-gender-equality-regime-hint-1}
# Hint 1 of 4
# After the filter() step, pipe into ggplot().
# Inside aes(), map the outcome variable to x
# and the explanatory variable to fill.
# Apply the three rules to decide which is which.
```

```{r exercise-create-proportions-gender-equality-regime-hint-2}
# Hint 2 of 4
# x = gender_equal3 (the outcome)
# fill = regime_type3 (the explanatory variable)
# Add geom_bar() with position = "dodge" and
# aes(y = after_stat(prop), group = regime_type3).
```

```{r exercise-create-proportions-gender-equality-regime-hint-3}
# Hint 3 of 4
# After geom_bar(), add labs() with the title,
# x = NULL, y = "Proportion", and the caption.
# Then add scale_fill_discrete() with name and labels.
```

```{r exercise-create-proportions-gender-equality-regime-hint-4}
# Hint 4 of 4
# scale_fill_discrete(name = "Regime Type",
#                     labels = c("Dictatorship",
#                                "Parliamentary\nDemocracy",
#                                "Presidential\nDemocracy"))
# List labels in the same order as levels(world$regime_type3).
```

```{r exercise-create-proportions-gender-equality-regime-check}
grade_this({
  code <- paste(.user_code, collapse = "\n")
  code_no_comments <- gsub("#.*", "", code)

  if (!inherits(.result, "ggplot")) {
    fail("Your code should return a ggplot object. If you're seeing an error rather than a plot, the most common causes are: a missing + at the end of a line, unmatched parentheses, or a label string that is missing its quotation marks.")
  }
  if (!grepl("filter\\s*\\(", code_no_comments)) {
    fail("Keep the filter() step so the plot removes missing values before plotting. Use the hints if needed, then click \"Submit Answer\" again.")
  }
  if (!grepl("!\\s*is\\.na\\s*\\(\\s*gender_equal3\\s*\\)", code_no_comments)) {
    fail("Use !is.na(gender_equal3) inside filter().")
  }
  if (!grepl("!\\s*is\\.na\\s*\\(\\s*regime_type3\\s*\\)", code_no_comments)) {
    fail("Use !is.na(regime_type3) inside filter().")
  }
  if (!grepl("x\\s*=\\s*gender_equal3", code_no_comments)) {
    fail("Map gender_equal3 to the x-axis inside aes(). Use the hints if needed, then click \"Submit Answer\" again.")
  }
  if (!grepl("fill\\s*=\\s*regime_type3", code_no_comments)) {
    fail("Map regime_type3 to the fill aesthetic inside aes().")
  }
  if (!grepl("geom_bar\\s*\\(", code_no_comments)) {
    fail("Use geom_bar() to create the bar plot.")
  }
  if (!grepl("position\\s*=\\s*['\"]dodge['\"]", code_no_comments)) {
    fail('Add position = "dodge" inside geom_bar() to place bars side-by-side.')
  }
  if (!grepl("y\\s*=\\s*after_stat\\s*\\(\\s*prop\\s*\\)", code_no_comments)) {
    fail("Set y = after_stat(prop) inside the aes() call within geom_bar() to plot proportions.")
  }
  if (!grepl("group\\s*=\\s*regime_type3", code_no_comments)) {
    fail("Set group = regime_type3 inside the aes() call within geom_bar().")
  }
  if (!grepl("labs\\s*\\(", code_no_comments)) {
    fail("Add a labs() layer to set the title and axis labels.")
  }
  if (!grepl("title\\s*=\\s*['\"]Gender Equality['\"]", code_no_comments)) {
    fail('Set title = "Gender Equality" inside labs().')
  }
  if (!grepl("x\\s*=\\s*NULL", code_no_comments)) {
    fail("Set x = NULL inside labs() to omit the x-axis label.")
  }
  if (!grepl("y\\s*=\\s*['\"]Proportion['\"]", code_no_comments)) {
    fail('Set y = "Proportion" inside labs().')
  }
  if (!grepl("caption\\s*=\\s*['\"]Country Data from the United Nations['\"]", code_no_comments)) {
    fail('Add caption = "Country Data from the United Nations" inside labs().')
  }
  if (!grepl("scale_fill_discrete\\s*\\(", code_no_comments)) {
    fail("Add scale_fill_discrete() to label the legend.")
  }
  if (!grepl("name\\s*=\\s*['\"]Regime Type['\"]", code_no_comments)) {
    fail('Set name = "Regime Type" inside scale_fill_discrete().')
  }
  code_all_labels <- paste(.user_code, collapse = "\n")
  if (!grepl("Dictatorship", code_all_labels)) {
    fail('Include "Dictatorship" as the first label in scale_fill_discrete().')
  }
  label_order_ok <- grepl(
    "Dictatorship.*Parliamentary.*Democracy.*Presidential.*Democracy",
    gsub("\\s+", " ", gsub("\\\\n", " ", code_all_labels))
  )
  if (!label_order_ok) {
    fail('The three labels must appear in this order: "Dictatorship", "Parliamentary\\nDemocracy", "Presidential\\nDemocracy" --- matching the stored factor level order shown by levels(world$regime_type3).')
  }

  pass("Your plot is correct. Compare the bar heights across regime types before reading the interpretation below.")
})
```

The completed plot presents the same information as the gender-equality cross-tab, but in visual form. We can quickly see that dictatorships have a lower proportion of countries with high gender equality than either democratic regime type, consistent with the hypothesis.

## Check Your Understanding

Now let’s step back and practice the full workflow with a new example. You do not need to write any code. Focus on the logic: identify the outcome and explanatory variables, read the column proportions, and decide whether the pattern is consistent with the hypothesis.

Suppose we are interested in the following question: Are countries in military alliances less likely to experience interstate conflict?

Alliances may deter conflict by raising the expected costs of attacking an alliance member. If potential challengers expect an alliance partner to help defend a country, they may be less likely to initiate a conflict.

::: {.hypothesis}
In a comparison of [countries], those that are [members of a military alliance] will be less likely to [experience an interstate conflict] than will those that are [not members of a military alliance].
:::

Suppose we summarize the relationship with the following cross-tab. The entries in each cell show the count, with the column proportion in parentheses.

| Conflict outcome | Not in alliance | In alliance |
|---|---:|---:|
| No interstate conflict | 70 (0.700) | 170 (0.850) |
| Interstate conflict | 30 (0.300) | 30 (0.150) |
| Column total | 100 | 200 |

**Apply: Check your understanding**

Read the scenario and cross-tab above carefully, then apply the three rules and the hypothesis-evaluation logic from this tutorial.

```{r cyu-alliance-crosstab-setup, echo=FALSE}
question("Which setup follows the rules we used for organizing this cross-tab?",
  correct = "Right. This setup follows all three rules.",
  answer("Conflict outcome should be in the rows, alliance status should be in the columns, and the table should use column proportions",
         correct = TRUE,
         message = "Yes. Conflict outcome is the outcome variable, alliance status is the explanatory variable, and column proportions let us compare outcomes within each alliance category."),
  answer("Alliance status should be in the rows, conflict outcome should be in the columns, and the table should use row proportions",
         message = "This reverses the structure we used in the tutorial. The outcome variable should go in the rows, and the explanatory variable should go in the columns."),
  answer("Conflict outcome should be in the rows, alliance status should be in the columns, and the table should use total proportions",
         message = "The row and column placement is right, but total proportions are not the best choice for comparing outcomes across alliance categories."),
  answer("The variables can go anywhere because the same counts appear in the table either way",
         message = "The counts may be the same, but the organization affects how easily we can interpret the relationship. We use a consistent structure to make comparisons clearer."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  try_again = "Incorrect. Try again."
)
```

```{r cyu-alliance-expected-pattern, echo=FALSE}
question("To evaluate whether the pattern is consistent with the hypothesis, which comparison in the cross-tab is the relevant one?",
  correct = "Right. The hypothesis is about whether conflict is less common among alliance members, so we compare the interstate conflict row proportion across the two columns.",
  answer("The proportion experiencing interstate conflict among alliance members versus the proportion experiencing interstate conflict among non-alliance members",
         correct = TRUE,
         message = "The hypothesis predicts that conflict is less common among alliance members. To evaluate it, we look at the Interstate conflict row and compare the column proportions across the two alliance categories. That is the comparison that tells us whether alliance membership is associated with lower conflict rates."),
  answer("The proportion of all conflict cases that belong to alliance members versus non-members",
         message = "That would be a row proportion --- the share of conflict cases in each alliance category. Column proportions tell us how common conflict is within each alliance category, which is what the hypothesis predicts."),
  answer("The total number of countries in each alliance category",
         message = "Column totals tell us how many countries fall in each alliance group, but not whether conflict rates differ between them. We need the proportions within each column to evaluate the hypothesis."),
  answer("The no-conflict count across the entire table",
         message = "A single count from the whole table cannot tell us whether conflict rates differ between alliance members and non-members. We need to compare proportions within each column."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  try_again = "Hint: The hypothesis predicts something about one specific outcome. Which row captures that outcome, and which proportions let you compare it across the alliance categories?"
)
```

```{r cyu-alliance-interpret-cell, echo=FALSE}
question("What does the entry 30 (0.150) in the Interstate conflict row and In alliance column mean?",
  correct = "Right. The proportion 0.150 is calculated within the In alliance column.",
  answer("Of the 200 countries in alliances, 30 experienced interstate conflict, or 15.0%",
         correct = TRUE,
         message = "Because these are column proportions, 0.150 is calculated out of the 200-country alliance column total --- not out of all 300 countries or all 30 conflict cases. The count 30 and the proportion 0.150 both describe the same alliance column."),
  answer("Of the 30 countries that experienced interstate conflict, 15.0% were in alliances",
         message = "This reverses the direction. The proportion is calculated within the In alliance column, not within the Interstate conflict row."),
  answer("Of all 300 countries in the table, 15.0% were in alliances and experienced interstate conflict",
         message = "That would be a total proportion. The table reports column proportions."),
  answer("The table shows that alliance membership caused 15.0% of countries to experience conflict",
         message = "A cross-tab describes the pattern in the data. It does not by itself establish causation."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  try_again = "Incorrect. Try again."
)
```

```{r cyu-alliance-consistent-hypothesis, echo=FALSE}
question("Which statement best describes whether the pattern is consistent with the hypothesis?",
  correct = "Right. The conflict proportion is lower among alliance members, consistent with the hypothesis.",
  answer("The pattern is consistent with the hypothesis because the proportion experiencing interstate conflict is lower among countries in alliances than among countries not in alliances",
         correct = TRUE,
         message = "Yes. The conflict proportion is 0.150 for countries in alliances and 0.300 for countries not in alliances."),
  answer("The pattern is not consistent with the hypothesis because the same number of countries in each column experienced interstate conflict",
         message = "The counts are both 30, but the column totals are different. That is why we compare column proportions rather than counts."),
  answer("The pattern is not consistent with the hypothesis because most countries in both columns avoided interstate conflict",
         message = "The hypothesis is about whether conflict is less common among alliance members, not whether conflict is rare in both groups."),
  answer("The cross-tab proves that alliances prevent interstate conflict",
         message = "The cross-tab shows a pattern that is consistent with the hypothesis, but it does not prove that alliances caused the difference."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  try_again = "Incorrect. Try again."
)
```



## The Takeaways

This tutorial introduced [**bivariate description**]{.important-text} of two categorical variables --- examining whether the distribution of a categorical outcome differs across the categories of a categorical explanatory variable. The basic move is simple but powerful --- compare the distribution of an outcome across categories of an explanatory variable and ask whether the pattern matches what a hypothesis predicts. Tutorials 8 and 9 will apply the same logic to other combinations of variable types, but the underlying question is always the same: does the pattern in the data point in the direction the hypothesis predicts?

- [**Cross-tabulations**]{.important-text} are the primary tool for describing the relationship between two categorical variables. A cross-tab shows how the distribution of an outcome variable differs across the categories of an explanatory variable. Three rules govern how we set them up and read them: the outcome variable goes in the rows, the explanatory variable goes in the columns, and we use column proportions. Column proportions show the share of cases within each column that fall in each row. To evaluate whether a pattern is consistent with a hypothesis, find the relevant outcome row and compare the column proportions across the categories of the explanatory variable. In R, `CrossTable()` from the [descr]{.package-name} package produces cross-tabs; use `prop.c = TRUE` for column proportions and set `prop.r`, `prop.t`, and `prop.chisq` to `FALSE` to keep the table clean.

- [**Grouped bar plots**]{.important-text} present the same information visually. The outcome variable goes on the x-axis. The explanatory variable maps to the `fill` argument inside `aes()`, which colors the bars by group. Use `position = "dodge"` inside `geom_bar()` to place the fill categories side by side. When the groups being compared have different sizes, plot proportions rather than counts by adding `aes(y = after_stat(prop), group = fill_variable)` inside `geom_bar()`. If a categorical variable is stored as a number, wrap it in `factor()` so `ggplot()` treats its values as discrete. Use `scale_x_discrete()` to relabel x-axis tick marks and `scale_fill_discrete()` to relabel the legend, including a `name` for the legend title and `labels = c(...)` for the category names. Use `\n` inside a label string to insert a line break when labels are long. Always filter missing values in the variables used in the plot before plotting. This keeps missing categories out of the display and avoids ggplot dropping observations or producing warnings. Use `filter(!is.na(variable))` for a single variable or `filter(!is.na(var1), !is.na(var2))` when both variables have potential missing values.

The cross-tab and the grouped bar plot are complementary: the table gives exact numbers, the plot makes the pattern easier to see at a glance. For evaluating a hypothesis, you usually want both. A cross-tab can show whether the observed pattern is *consistent with* a hypothesis, but it does not tell us whether the difference is larger than we would expect by chance --- that requires a formal hypothesis test, covered in Tutorial 12.

In your own R scripts, call `library(descr)`, `library(ggplot2)`, and `library(dplyr)` once at the top of the script before using any functions from those packages. Loading a package once per session is enough --- you do not need to repeat the `library()` call every time you use a function from it.

Tutorial 8 applies the same logic to a different combination of variable types --- one categorical and one interval --- where the main tool shifts from column proportions to a comparison of group means. The question is the same: does the pattern in the data point in the direction the hypothesis predicts?
    
## Submit Your Work

When you have completed all exercises in this tutorial, generate your submission code using the button below. Copy the full code that appears and paste it into the Canvas submission box for this tutorial. The code can be long --- make sure you have copied all of it before submitting.

```{r encode-logic-07, context="server"}
learnrhash::encoder_logic()
```

```{r encode-ui-07, echo=FALSE}
learnrhash::encoder_ui(
  ui_before = shiny::div(
    shiny::p("Click the button below to generate your submission code.
              Copy the full code and paste it into Canvas.")
  )
)
```

```{r tutorial-success-banner-07, results="asis", echo=FALSE}
cat('
<p class="tutorial-complete"
   role="status"
   aria-label="Nicely done! You have completed this tutorial and can now examine relationships between categorical variables using cross-tabs and grouped bar plots.">
  <span aria-hidden="true">🎉</span>
  Nicely done! You can now examine relationships between categorical variables using cross-tabs and grouped bar plots.
</p>
')
```
