---
title: "Univariate Description: Interval Variables"
tutorial:
  id: "06-Interval-Description"
output: 
  QPATutorialsCourse::qpa_tutorial:
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
description: >
  Learn how to describe an interval variable: central tendency, dispersion, histogram and box plot.
---


## Overview {#main-content}

This tutorial completes the three-tutorial sequence on [**univariate description**]{.important-text} --- the practice of examining one variable at a time. In Tutorial 4, you described **nominal variables**, where categories are different but not ranked. In Tutorial 5, you described **ordinal variables**, where categories have a meaningful order. In this tutorial, you will describe **interval variables**, where the values are ordered and the distances between values are meaningful too.

That last step gives you more room to work. Once equal differences mean something, you can use the mean, standard deviation, range, percentiles, histograms, and box plots to describe a distribution. The familiar questions are still the same: what is typical, how spread out are the values, and what overall shape does the distribution take? The tools are richer now, but the goal is still to help a reader understand one variable clearly.

In this tutorial, you will learn:

* How to determine and describe the central tendency of an interval level variable
* How to use the `mean()` function
* How to summarize the dispersion of an interval variable using `range()`, `quantile()`, and `sd()`
* How to create histograms and box plots using `ggplot()` with `geom_histogram()` and `geom_boxplot()`

```{r setup, include=FALSE}
library(learnr)
library(learnrhash)
library(tidyverse)
library(knitr)
library(poliscidata)
library(gradethis)
tutorial_options(exercise.checker = gradethis::grade_learnr,
                 exercise.completion = FALSE)
if (max(counties$employ_pop_ratio_25_64, na.rm = TRUE) < 10)
  counties$employ_pop_ratio_25_64 <- counties$employ_pop_ratio_25_64 * 100
counties$dem2p_percent <- counties$dem2p_vote_share * 100
if (max(counties$wage_growth, na.rm = TRUE) < 2)
  counties$wage_growth <- counties$wage_growth * 100
```

## Univariate Description

In Tutorial 4, you described nominal variables using frequency and proportion tables, the mode, and bar plots. In Tutorial 5, you built on those tools for ordinal variables, adding the median because ordinal categories have a meaningful order. In this tutorial, you will extend the same logic to [**interval variables**]{.important-text}. Because interval variables are ordered *and* have meaningful distances between values, they allow for an even richer set of descriptive tools.

**Practice: Identify interval-level variables**

Think about which variables have values that are not only ordered but also have equal distances between them throughout the scale.

```{r concepts-quiz-interval-level, echo=FALSE}
question("Which of the following variables are interval level measures? Select all that apply.",
  correct = "Right. All three are interval-level variables.",
  answer("The number of terrorist acts in the Middle East in a given year", correct = TRUE,
         message = "A count of terrorist acts is interval-level. The values are ordered and the differences between them are equal --- five more acts is five more acts anywhere on the scale. Neither ideology scales nor ordinal rankings have this property."),
  answer("The poverty rate in a country", correct = TRUE,
         message = "A rate measured as a percentage is interval-level. The values are ordered and a one-point difference means the same thing anywhere on the scale."),
  answer("The number of COVID-19 cases daily worldwide", correct = TRUE,
         message = "A daily count is interval-level. The values are ordered and equal differences carry equal meaning --- 1,000 more cases is 1,000 more cases regardless of the baseline."),
  answer("Individual ideology measured on a scale of 1 to 7 from very liberal to very conservative",
         message = "Ideology on a 1--7 scale is ordinal. The values are rank-ordered from liberal to conservative, but the differences between adjacent points are not necessarily equal --- a move from 3 to 4 may not represent the same ideological shift as a move from 6 to 7."),
  answer("The Polity democracy score measured on a scale from -10 (autocracy) to +10 (democracy)",
         message = "The Polity score is technically ordinal because the differences between values are not guaranteed to be equal. However, as discussed in Tutorial 3, analysts frequently treat it as interval-level in practice --- a convention worth knowing."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  try_again = "Incorrect. Try again."
)
```

## Central Tendency

For interval variables, you can describe [**central tendency**]{.important-text} using the mode, median, *and* mean. That does not mean all three are equally useful. As you will see, the mode is usually not very informative for interval data, and the mean can be strongly affected by outliers, or extreme values. The median and mean will be the primary tools.

You have already encountered the median. The median is the value in the middle when you list all values in rank order for all cases. Thus it is the value at the 50th percentile: half the cases have values at or below the median, and half have values at or above the median.


You are likely very familiar with the mean, or average, value. The arithmetic [**mean**]{.important-text} is computed by adding up the values for all cases and dividing by the number of cases:

$$\bar{Y} = \frac{Y_1 + Y_2  + ... + Y_n}{n} = \frac{1}{n} \Sigma_{i=1}^n Y_i$$

Can you tell why the mean is not an appropriate measure of central tendency for nominal and ordinal variables? You cannot add up categories like married, single, divorced, and widowed. Even if they are coded numerically, what would the mean of marital status *mean*? It is also an inappropriate measure of central tendency for an ordinal variable. Although you may be able to add up the values of an ordinal variable, because the differences between them are not equal, the mean is misleading.
 
 
The mean is the most commonly reported measure of central tendency, but it has one important weakness: it is sensitive to outliers. Suppose you wanted to describe the income of patrons at a bar in State College. On most nights, the median and mean income would probably be fairly close. But what happens if Elon Musk walks in? The median income barely moves, because one person cannot change the middle case very much. The mean, however, would increase dramatically. For this reason, the median is said to be **robust** to outliers, but the mean is not.

### The mode

For nominal and ordinal variables, you identified the mode from a frequency table. Try the same approach with the interval variable **wage_growth** in the **counties** data frame. This variable measures the percentage rate of increase in wages for all employed persons in a county over the four years preceding the 2016 presidential election. Run the code below to see what happens.

**Run and observe: Find the mode of wage_growth**

```{r Freq_Prop_table_for_wage_growth, exercise=TRUE, exercise.cap="Run and observe: Why the mode is not useful for interval data"}
library(dplyr)
counties %>%
  filter(!is.na(wage_growth)) %>%
  group_by(wage_growth) %>%
  summarise(freq = n()) %>%
  mutate(prop = freq / sum(freq))
```

Can you see the problem? The values of **wage_growth** are so precisely measured that almost no two counties have the exact same wage growth rate --- each value appears just once, so there is no single most common value. The mode tells you nothing useful here. This is the typical situation with interval data. **For this reason, the mode is almost never reported for an interval variable.** Instead, focus on the median and mean.

Once a package is loaded with `library()`, all of its functions are available for the rest of your R session --- you do not need to load it again each time you use a function from it. Since [dplyr]{.package-name} was loaded in the exercise above, you can assume it is already loaded for all subsequent exercises in this tutorial.

### The median

The frequency table showed why the mode is not helpful for wage growth. A better way to describe the typical county is the median, which finds the wage-growth value at the middle of the ordered distribution.

**Practice: Compute the median of wage_growth**

Find the value at the 50th percentile of the wage growth distribution. Three hints are available. Click "Submit Answer" to check your work.

```{r compute_median_wage_growth, exercise=TRUE, exercise.cap="Practice: Compute the median of wage_growth"}

```

```{r compute_median_wage_growth-hint-1}
# Hint 1 of 3
# To describe the typical county's wage growth, you need
# the value that divides the distribution in half ---
# the value at the 50th percentile, with half the counties
# above it and half below. What function computes
# this for a numeric variable?
```

```{r compute_median_wage_growth-hint-2}
# Hint 2 of 3
# The function is median(). Access the variable using $
# and include na.rm = TRUE to handle any missing values.
```

```{r compute_median_wage_growth-hint-3}
# Hint 3 of 3
# The function name is median().
# Pass it the variable using the $ operator:
# counties$wage_growth
# Include na.rm = TRUE as a second argument.
```

```{r compute_median_wage_growth-check}
grade_this({
  code <- paste(.user_code, collapse = "\n")
  code_no_comments <- gsub("#.*", "", code)

  if (!grepl("median\\s*\\(", code_no_comments)) {
    fail("Use the median() function.")
  }
  if (!grepl("wage_growth", code_no_comments)) {
    fail("Apply median() to the wage_growth variable.")
  }
  if (is.null(.result)) {
    fail("Make sure your code produces a result the grader can check.")
  }
  if (!is.numeric(.result) || length(.result) != 1) {
    fail("Your result should be a single number: the median of wage_growth.")
  }
  if (is.na(.result)) {
    fail("Your result is NA. Make sure missing values are removed before computing the median.")
  }
  if (!isTRUE(all.equal(as.numeric(.result),
                        as.numeric(median(counties$wage_growth, na.rm = TRUE))))) {
    fail("Check that you are computing the median of wage_growth in the counties data frame.")
  }
  pass(paste0("Right! The median of wage_growth is ",
              round(median(counties$wage_growth, na.rm = TRUE), 2),
              "%."))
})
```

The median value of wage growth is 4.70%, which means that after ranking all counties' wage growth rates, the value at the 50th percentile was just under 5% growth. Half the counties experienced wage growth at or below 4.7%.

### The mean

We can calculate the mean using the `mean()` function. If there are any missing values in the data, include `na.rm = TRUE` to exclude them. In fact, it is good practice to always include `na.rm = TRUE` when calling `mean()`.

Now calculate the mean so you can compare it to the median. When the mean and median are close together, the distribution is usually fairly balanced around the center. When they are far apart, that difference is often a clue that outliers or skew are shaping the distribution.

**Practice: Calculate the mean of wage_growth**

Compute the arithmetic average of **wage_growth**. Three hints are available. Click "Submit Answer" to check your work.

```{r compute_mean_wage_growth, exercise=TRUE, exercise.cap="Practice: Calculate the mean of wage_growth"}

```

```{r compute_mean_wage_growth-hint-1}
# Hint 1 of 3
# You need the arithmetic average of wage growth ---
# the sum of all county values divided by the count
# of counties. This is one of the two main measures
# of central tendency for interval variables.
# What function in R computes this?
```

```{r compute_mean_wage_growth-hint-2}
# Hint 2 of 3
# The function is mean(). Access the variable using $
# and include na.rm = TRUE to handle any missing values.
```

```{r compute_mean_wage_growth-hint-3}
# Hint 3 of 3
# The function name is mean().
# Pass it the variable using the $ operator:
# counties$wage_growth
# Include na.rm = TRUE as a second argument.
```

```{r compute_mean_wage_growth-check}
grade_this({
  code <- paste(.user_code, collapse = "\n")
  code_no_comments <- gsub("#.*", "", code)

  if (!grepl("mean\\s*\\(", code_no_comments)) {
    fail("Use the mean() function.")
  }
  if (!grepl("wage_growth", code_no_comments)) {
    fail("Apply mean() to the wage_growth variable.")
  }
  if (is.null(.result)) {
    fail("Make sure your code produces a result the grader can check.")
  }
  if (!is.numeric(.result) || length(.result) != 1) {
    fail("Your result should be a single number: the mean of wage_growth.")
  }
  if (is.na(.result)) {
    fail("Your result is NA. Make sure missing values are removed before computing the mean.")
  }
  if (!isTRUE(all.equal(as.numeric(.result),
                        as.numeric(mean(counties$wage_growth, na.rm = TRUE))))) {
    fail("Check that you are computing the mean of wage_growth in the counties data frame.")
  }
  pass(paste0("Good work! The mean of wage_growth is ",
              round(mean(counties$wage_growth, na.rm = TRUE), 2),
              "%."))
})
```

When the mean and median are close, the distribution may be fairly balanced around the center --- though you should always check a histogram to confirm.

Now try the same comparison with a more politically familiar variable. **dem2p_percent** is the percent of the two-party vote cast for the Democrat in a county in 2016. Calculate the median and mean on separate lines, then compare them. The gap between those two numbers will give you an early clue about the shape of the county-level distribution of Democratic vote percentages.

**Practice: Compute the median and mean of dem2p_percent**

Compute both statistics on separate lines, then compare them --- the gap between them is an early clue about the shape of the distribution. Two hints are available. Click "Submit Answer" to check your work.

```{r compute_mean_median_dem2p_percent, exercise=TRUE, exercise.lines=4, exercise.cap="Practice: Compute the median and mean of dem2p_percent"}

```

```{r compute_mean_median_dem2p_percent-hint-1}
# Hint 1 of 2
# You need two lines of code, one for each statistic.
# Use median() and mean(), each applied to
# counties$dem2p_percent with na.rm = TRUE.
```

```{r compute_mean_median_dem2p_percent-hint-2}
# Hint 2 of 2
# The two function names are median() and mean().
# Each takes counties$dem2p_percent as its first
# argument and na.rm = TRUE as its second.
# Put each on its own line.
```

```{r compute_mean_median_dem2p_percent-check}
grade_this({
  code_lines <- gsub("#.*", "", .user_code)
  code_no_comments <- paste(code_lines, collapse = "\n")

  if (!grepl("median\\s*\\(", code_no_comments)) {
    fail("Include a call to median() for dem2p_percent.")
  }
  if (!grepl("mean\\s*\\(", code_no_comments)) {
    fail("Include a call to mean() for dem2p_percent.")
  }
  if (!grepl("median\\s*\\([^)]*dem2p_percent", code_no_comments)) {
    fail("Your median() call should use dem2p_percent: median(counties$dem2p_percent, na.rm = TRUE).")
  }
  if (!grepl("mean\\s*\\([^)]*dem2p_percent", code_no_comments)) {
    fail("Your mean() call should use dem2p_percent: mean(counties$dem2p_percent, na.rm = TRUE).")
  }

  expected_median <- median(counties$dem2p_percent, na.rm = TRUE)
  expected_mean <- mean(counties$dem2p_percent, na.rm = TRUE)

  if (is.null(.result)) {
    fail("Make sure your code produces a result the grader can check.")
  }
  result_val <- as.numeric(.result)
  if (!is.numeric(.result) || length(result_val) != 1 || is.na(result_val)) {
    fail("Your final line should return either the median or the mean of dem2p_percent.")
  }
  if (!isTRUE(all.equal(result_val, expected_mean)) &&
      !isTRUE(all.equal(result_val, expected_median))) {
    fail("Put median() on one line and mean() on the other. Make sure both are computed for dem2p_percent.")
  }

  pass(paste0("Correct! The median Democratic vote percentage was ",
              round(expected_median, 1),
              "% and the mean was ",
              round(expected_mean, 1),
              "%."))})
```

The median percent of the two-party vote cast for Clinton was just under 30% --- half the counties cast 30% or fewer of their two-party votes for her. The mean of about 33% is higher than the median, suggesting that some counties with very high Democratic vote percentages are pulling the mean upward. The median is much less affected by those extreme values because it depends on the ordering of observations rather than how far an extreme value lies from the rest --- but the mean moves whenever a value changes. When the mean exceeds the median, the distribution is typically right-skewed: a long right tail of high values pulls the mean up while the median stays closer to where most observations are.

## Dispersion

Once you know where the center of a distribution is, the next question is how spread out the values are. For interval variables, you have several ways to describe [**dispersion**]{.important-text}. You will consider:

1. [**Range**]{.important-text}
2. [**Percentiles**]{.important-text}, especially the [**interquartile range**]{.important-text}
3. [**Variance**]{.important-text} and [**standard deviation**]{.important-text}


## Range 

The [**range**]{.important-text} of values a variable takes is given by the smallest (minimum) value and the largest (maximum) value. As such, it provides bookends on the spread of the values of the variable across cases. It is useful to report the range of the data in the sample because the sample range may be different from the theoretical range of the data. It can also help identify the presence of outliers.

### The range(), min(), and max() functions

The R function `range()` directly calculates both the minimum and maximum values in the data. The function `min()` calculates the minimum value of a variable, and the `max()` function calculates the maximum value. Each function is passed the name of the variable being analyzed. If there are missing values, include the `na.rm` argument and set it equal to `TRUE` so that R ignores those values when computing the statistic.

Let's calculate the range of **wage_growth** to see just how much variation there was across counties. Since wages can fall, wage growth need not be positive at all.

**Practice: Compute the range, min, and max of wage_growth**

See how far the distribution stretches from its most extreme values. Put each function call on its own line. Three hints are available. Click "Submit Answer" to check your work.

```{r compute_range_min_max_wage_growth, exercise=TRUE, exercise.cap="Practice: Compute the range, minimum, and maximum of wage_growth"}

```

```{r compute_range_min_max_wage_growth-hint-1}
# Hint 1 of 3
# You need three separate function calls on the same
# variable: one that returns both endpoints together,
# and two that return each endpoint separately.
# Think about what three functions from Tutorials 1
# and 2 give you the minimum, the maximum, and
# both endpoints at once.
```

```{r compute_range_min_max_wage_growth-hint-2}
# Hint 2 of 3
# The three functions are range(), min(), and max().
# Apply each to counties$wage_growth with na.rm = TRUE,
# one per line.
```

```{r compute_range_min_max_wage_growth-hint-3}
# Hint 3 of 3
# The three function names are range(), min(), and max().
# Each takes counties$wage_growth as its first argument
# and na.rm = TRUE as its second.
# Put each function call on its own line.
```

```{r compute_range_min_max_wage_growth-check}
grade_this({
  code_lines <- gsub("#.*", "", .user_code)
  code_no_comments <- paste(code_lines, collapse = "\n")

  if (!grepl("range\\s*\\(", code_no_comments)) {
    fail("Include a call to range() for wage_growth.")
  }
  if (!grepl("min\\s*\\(", code_no_comments)) {
    fail("Include a call to min() for wage_growth.")
  }
  if (!grepl("max\\s*\\(", code_no_comments)) {
    fail("Include a call to max() for wage_growth.")
  }
  if (!grepl("range\\s*\\([^)]*wage_growth", code_no_comments)) {
    fail("Your range() call should use wage_growth: range(counties$wage_growth, na.rm = TRUE).")
  }
  if (!grepl("min\\s*\\([^)]*wage_growth", code_no_comments)) {
    fail("Your min() call should use wage_growth: min(counties$wage_growth, na.rm = TRUE).")
  }
  if (!grepl("max\\s*\\([^)]*wage_growth", code_no_comments)) {
    fail("Your max() call should use wage_growth: max(counties$wage_growth, na.rm = TRUE).")
  }

  expected_range <- range(counties$wage_growth, na.rm = TRUE)
  expected_min <- min(counties$wage_growth, na.rm = TRUE)
  expected_max <- max(counties$wage_growth, na.rm = TRUE)

  if (is.null(.result)) {
    fail("Make sure your code produces a result the grader can check.")
  }
  result_val <- as.numeric(.result)

  if (any(is.na(result_val)) ||
      !(isTRUE(all.equal(result_val, expected_range)) ||
        isTRUE(all.equal(result_val, expected_min)) ||
        isTRUE(all.equal(result_val, expected_max)))) {
    fail("Put range(), min(), and max() on separate lines, all applied to wage_growth with missing values removed.")
  }

  pass(paste0("Exactly! Wage growth ranged from ",
              round(min(counties$wage_growth, na.rm = TRUE), 1),
              "% to ",
              round(max(counties$wage_growth, na.rm = TRUE), 1),
              "% --- an extraordinarily wide spread, with some counties seeing wages fall sharply and others growing rapidly."))
})
```

Those extreme values are worth understanding. You might expect outlier counties to be visible in a histogram, but you probably will not see them. With over 3,000 counties in the data, a handful of outlier counties contribute bars so short --- perhaps one or two counties per bin --- that they are nearly invisible against the bars near the center, which each represent hundreds of counties. The x-axis stretches to include the extremes, but the bars themselves disappear into the baseline. This is one reason why computing the range is good practice even when you have a histogram: the two tools can tell different stories about the same data.

## Percentiles

A [**percentile**]{.important-text} is the value of a variable that has a certain percentage of observations below it. In other words, the nth percentile is the value that cuts off the first n percent of the data when all values are sorted from least to greatest. You have already worked with the 50th percentile: it is the median. The 90th percentile is the value that separates the bottom 90% of values from the top 10%.

Percentiles can be used to answer questions such as:

::: {.research-example}
Among counties in the bottom quarter of the life-expectancy distribution, what is the highest life expectancy?
:::

To answer this question, we'd find the 25th percentile of life expectancy across counties. This is the life expectancy value that separates the bottom 25% of observations from the top 75%.

::: {.research-example}
What are the percentages of working-age adults who are employed that cover the middle half of all counties?
:::

To answer this question, find the 75th percentile of the proportion of working-age adults who are employed and the 25th percentile. These are the two values that determine the upper and lower bounds for the middle 50% of the data.


The distance between the 25th and 75th percentiles is known as the [**interquartile range**]{.important-text} or [**IQR**]{.important-text}. The quartiles divide the ordered observations into four groups, each containing approximately 25% of the cases. The IQR spans the middle 50% of the observations. Its advantage over the full range as a measure of spread is that extreme values do not affect it.


### The quantile() function

You can calculate percentiles with the `quantile()` function. By default, passing the variable name to the function will return the 0th, 25th, 50th,  75th, and 100th percentiles. (Be sure to include `na.rm = TRUE` if there are missing values for any case.) 

**Run and observe: Calculate the percentiles for wage_growth**
```{r calculate_percentile_wage_growth, exercise=TRUE,exercise.cap="Run and observe: Default quantiles for wage_growth"}
quantile(counties$wage_growth, na.rm = TRUE)
```

Note that the 0% is the minimum and the 100% is the maximum value, giving the range of the variable's values.

Focus on the interquartile range, the values at the 25th and 75th percentile, which demarcate the middle half of the data. The IQR tells you that the growth rate of wages for half the counties was between 2.5% and 6.5%. Notice how far removed from the middle the minimum and maximum values are --- consistent with the extreme range found earlier when computing `min()` and `max()`.

You can use the `quantile()` function to list additional percentiles by adding the `probs` argument inside the function. To list a specific set of percentiles set `probs = c(0.05, .1, .25, .5, .75, .9, .95)`.

**Run and observe: Calculate the additional percentiles for wage_growth**
```{r calculate_additional_percentile_wage_growth, exercise=TRUE,exercise.cap="Run and observe: Custom percentiles for wage_growth"}
quantile(counties$wage_growth, na.rm = TRUE, 
         probs = c(0.05, .1, .25, .5, .75, .9, .95))
```

The result shows that 90% of the cases fall between -2.60% and 9.92%. In other words, the bottom 5% of counties had wage growth of -2.60% or lower, meaning wages declined by at least about 2.6% in those counties, and the top 5% of counties experienced wage increases approaching 10% or higher.


Now use percentiles to describe a different kind of county economy measure: **employ_pop_ratio_25_64**, the percentage of working-age adults who are employed. The default quantiles will show the minimum, the maximum, the median, and the middle half of counties, giving you a compact picture of both center and spread.

**Practice: Calculate the percentiles for employ_pop_ratio_25_64**

The default quantiles will give you the minimum, the IQR boundaries, the median, and the maximum in one call --- a compact picture of both center and spread. Three hints are available. Click "Submit Answer" to check your work.

```{r calculate_percentile_employ_pop_ratio_25_64, exercise=TRUE, exercise.cap="Practice: Calculate percentiles for employ_pop_ratio_25_64"}

```

```{r calculate_percentile_employ_pop_ratio_25_64-hint-1}
# Hint 1 of 3
# You want to see the minimum, the 25th percentile,
# the median, the 75th percentile, and the maximum ---
# the five values that show where the distribution
# falls at even intervals. The function you need
# returns all five by default in a single call.
# What function produces this output?
```

```{r calculate_percentile_employ_pop_ratio_25_64-hint-2}
# Hint 2 of 3
# The function is quantile(). Access the variable using $
# and include na.rm = TRUE. No probs argument is needed
# for the default percentiles.
```

```{r calculate_percentile_employ_pop_ratio_25_64-hint-3}
# Hint 3 of 3
# The function name is quantile().
# Pass it counties$employ_pop_ratio_25_64 as the
# first argument and na.rm = TRUE as the second.
# No probs argument is needed for the default output.
```

```{r calculate_percentile_employ_pop_ratio_25_64-check}
grade_this({
  code <- paste(.user_code, collapse = "\n")
  code_no_comments <- gsub("#.*", "", code)

  if (!grepl("quantile\\s*\\(", code_no_comments)) {
    fail("Use the quantile() function.")
  }
  if (!grepl("employ_pop_ratio_25_64", code_no_comments)) {
    fail("Apply quantile() to employ_pop_ratio_25_64.")
  }

  expected <- quantile(counties$employ_pop_ratio_25_64, na.rm = TRUE)

  if (is.null(.result)) {
    fail("Make sure your code produces a result the grader can check.")
  }
  if (!is.numeric(.result) || length(.result) != length(expected)) {
    fail("Your result should show the default quantiles: minimum, 25th percentile, median, 75th percentile, and maximum.")
  }
  if (any(is.na(.result))) {
    fail("Your result includes NA. Make sure missing values are removed before computing quantiles.")
  }
  if (!isTRUE(all.equal(as.numeric(.result), as.numeric(expected), tolerance = 1e-8))) {
    fail("Check that you are computing the default quantiles for employ_pop_ratio_25_64.")
  }

  pass("Good work! The quantiles are correct.")
})
```

What is the range of the variable? 9.21% to 668.07%. How can this be? Whenever you see percents that are greater than 100 it is possible that there is a data error. Here, however, values above 100% are possible because the numerator, employment, is calculated based on the number of people working in a county, and the denominator, population aged 25-64, is calculated based on the population in that county. Percentages greater than 100 will occur if many workers commute to a different county than that in which they live. This happens in big cities like New York City and Los Angeles.

The IQR is 51.23% to 83.53%. This means that the middle half of counties have employment-to-population ratios between about 51% and 84% for adults ages 25--64.

You can also extract the median from the output. The middle value in the data is 66.61%. Half of counties have an employment-to-population ratio of 66.61% or lower for adults ages 25--64.

## Standard deviation (and variance)

The [**standard deviation**]{.important-text} describes the typical distance of values from the mean of the variable. The [**variance**]{.important-text} is the squared standard deviation.
 
The formula looks a little intimidating at first, but the idea is simple: measure how far each case is from the mean, average those distances in a way that keeps them from canceling out, and then put the result back into the original units.

The formula for the standard deviation, denoted by *s*, is
 
$$s=\sqrt{\frac{\Sigma_{i=1}^n({Y_i-\bar{Y}})^2}{n-1}}$$
 
Let's pick this apart by considering an example where the variable of interest is the height of presidents from Johnson through Biden. Below is a table listing the heights (in centimeters) of 10 selected presidents. Note that Bush II and Biden are the same height (182 cm) --- a useful reminder that tied values are perfectly possible and handled naturally by the formula. The table also presents the deviations in column three and deviations squared in column four. The last row sums the values in each column, giving you the building blocks to calculate the standard deviation.
 

<table id="pres-stats-table" aria-labelledby="pres-stats-caption"> <caption id="pres-stats-caption">Presidential heights and deviations from the mean height, used to illustrate variance and standard deviation.</caption> <thead> <tr> <th scope="col">President</th> <th scope="col">\(Y_i\)</th> <th scope="col">Deviation \((Y_i - \bar{Y})\)</th> <th scope="col">Deviation\(^2\) \((Y_i - \bar{Y})^2\)</th> </tr> </thead> <tbody> <tr><th scope="row">Trump</th><td>188</td><td>188 − 184.9 = 3.1</td><td>9.61</td></tr> <tr><th scope="row">Biden</th><td>182</td><td>182 − 184.9 = −2.9</td><td>8.41</td></tr> <tr><th scope="row">Obama</th><td>185</td><td>185 − 184.9 = 0.1</td><td>0.01</td></tr> <tr><th scope="row">Bush II</th><td>182</td><td>182 − 184.9 = −2.9</td><td>8.41</td></tr> <tr><th scope="row">Clinton</th><td>188</td><td>188 − 184.9 = 3.1</td><td>9.61</td></tr> <tr><th scope="row">Bush I</th><td>188</td><td>188 − 184.9 = 3.1</td><td>9.61</td></tr> <tr><th scope="row">Reagan</th><td>185</td><td>185 − 184.9 = 0.1</td><td>0.01</td></tr> <tr><th scope="row">Carter</th><td>177</td><td>177 − 184.9 = −7.9</td><td>62.41</td></tr> <tr><th scope="row">Ford</th><td>183</td><td>183 − 184.9 = −1.9</td><td>3.61</td></tr>  <tr><th scope="row">Johnson</th><td>191</td><td>191 − 184.9 = 6.1</td><td>37.21</td></tr> </tbody> <tfoot> <tr> <th scope="row">Summary</th> <td>\(\bar{Y}=\frac{1}{n}\sum_{i=1}^n Y_i=184.9\)</td> <td>\(\sum_{i=1}^n (Y_i-\bar{Y})=0\)</td> <td>\(\sum_{i=1}^n (Y_i-\bar{Y})^2=148.9\)</td> </tr> </tfoot> </table>

<br>

Steps to calculate the standard deviation:

1. Calculate the deviation --- how far above or below the mean each observation is: (Y~i~ − Ȳ) --- for each observation (column three). Observations below the mean have negative deviations, and observations above the mean have positive deviations. If the variable is not very dispersed, the deviations will be close to zero. Bigger deviations indicate more dispersion.
2. By definition, the average deviation will always be zero: deviations above the mean offset deviations below the mean. To avoid this problem, we square the deviations (column four).
3. Now we can calculate the variance by summing the squared deviations and dividing by n-1 rather than n. We divide by n-1 rather than n because we are estimating the population standard deviation from a sample; dividing by n-1 corrects for a small but systematic bias that would otherwise make the estimate too small. Here, n = 10, so we divide by 9: 148.9/9 = 16.54.
4. If we stop here, we have the variance. Taking the square root gets us back to the original units of the variable. This is called the standard deviation:

$$s=\sqrt{\frac{\Sigma_{i=1}^n({Y_i-\bar{Y}})^2}{n-1}}= \sqrt{\frac{148.9}{9}}=4.07$$

Note that *s* is always positive.

A larger standard deviation means values are more spread out from the mean; a smaller one means they cluster more closely around it. One important caution: **you cannot meaningfully compare standard deviations across different variables.** A standard deviation of 4.68 for `wage_growth` (measured in percentage points) and a standard deviation of 15.3 for `dem2p_percent` (also percentage points but on a completely different scale with a different mean) do not tell you which variable is "more dispersed" in any general sense. Standard deviations are only comparable within a variable or across samples of the same variable.

### The sd() function

To calculate the standard deviation, $s$, of a variable we use the `sd()` function, including `na.rm = TRUE` to remove any missing values.

**Run and observe: Calculate the standard deviation of wage_growth**

```{r calculate_sd, exercise=TRUE, exercise.cap="Run and observe: Calculating the standard deviation with sd()"}
sd(counties$wage_growth, na.rm = TRUE)
```

A standard deviation of 4.68 percentage points tells you that county wage-growth rates typically vary by about 4.68 percentage points around the mean of 4.33%. You will use the standard deviation extensively when discussing statistical inference in Tutorial 10, where the Empirical Rule --- connecting the standard deviation to the percentage of observations falling within specific ranges --- will be developed in full.


## The shape of the distribution

Assessing the [**shape of the distribution**]{.important-text} of an interval variable is most easily done after visualizing the distribution in a plot. The key questions are:

- Is the distribution [**symmetric**]{.important-text} --- does it look roughly the same on both sides of the center --- or [**skewed**]{.important-text} toward one tail?
- Is it [**unimodal**]{.important-text} (one peak), [**bimodal**]{.important-text} (two peaks), or something else?
- Are there [**outliers**]{.important-text} --- observations far from the rest of the distribution that may warrant investigation?

Not all variables will have symmetric distributions. Instead, they may be skewed. Histograms are said to be skewed to the side of the long tail. The two histograms below are clearly not symmetric. The image on the left is skewed to the right. The image on the right is skewed to the left.

![](images/SkewIllustration.png){width="70%" fig.alt="Two histograms shown side-by-side. Left panel: positively skewed (right-skew), with many low values near 0–3 and a long sparse tail extending right toward ~10. Right panel: negatively skewed (left-skew), with many high values near 15–20 and a long sparse tail extending left toward 0. Both use equal-width bins with semi-transparent blue bars."}

The histogram and box plot introduced in the next section are the primary tools for assessing shape.

## Visualizing the distribution of an interval variable

For the same reason a frequency table is not very helpful for interval variables, a bar plot is not the right visualization either. Instead, we use a [**histogram**]{.important-text}. A histogram groups observed values into **bins** that cover **ranges** of values, then uses the height of each bar to show how many observations fall in that range. Histograms can also be constructed to show proportions rather than counts. We can control the number of bins, control their width, or let R choose a default.


Histograms illustrate the central tendency of the variable and its dispersion, as well as the shape of the distribution.



###  The histogram

Generating a histogram with the `ggplot()` function requires a simple change to the basic code for a bar plot. Because we want a visual display for interval data, we replace `geom_bar()` with `geom_histogram()`. The basic structure is familiar, but the interpretation changes because each bar now represents a range of values rather than one category. Note that we do not wrap **wage_growth** in `factor()`, `factor()` is only needed when a variable is stored as numbers but should be treated as categorical. Interval variables are genuinely numeric, so no wrapping is needed.

As with proportion bar charts, filter out missing values before passing data to `ggplot()`. If you do not filter, ggplot will still produce the plot but will print a warning message in your output --- "Removed N rows containing non-finite values" --- which would appear in a knitted R Markdown document and looks unprofessional. Filtering first keeps your output clean. Always note in your text how many cases were excluded.

Inside `geom_histogram()`, two arguments control the appearance of the bars. `fill` sets the interior color of the bars, just as it did in the bar charts in Tutorials 4 and 5. `color` sets the border color --- adding white borders between bars makes individual bins visually separable, which is especially helpful when bins are narrow. ggplot chose 30 bins by default and may suggest you consider alternatives; we address that in the exercises that follow.

**Run and observe: Generate the histogram for wage_growth**

Like [dplyr]{.package-name}, [ggplot2]{.package-name} only needs to be loaded once per session --- all subsequent exercises in this tutorial can use it without calling `library(ggplot2)` again.

```{r generate_histogram_wage_growth, exercise=TRUE, warning=TRUE, exercise.cap="Run and observe: A basic histogram of wage_growth"}
library(ggplot2)
counties %>%
  filter(!is.na(wage_growth)) %>%
  ggplot(mapping = aes(x = wage_growth)) +
         geom_histogram(fill = "steelblue", color = "white") +
         labs(title = "Wage Growth in US Counties in 2016", 
              x = NULL, 
              y = "Frequency", 
              caption = "Source: Linn, Nagler, Zilinsky")
```



There are two options for specifying a different number of bins. You can set the number of bins by adding the `bins` argument to `geom_histogram()` and setting it equal to some number.


**Practice: Generate the histogram for wage_growth with a different number of bins**

Change the number of bins in the plot to 50. One hint is available. Click "Submit Answer" to check your work.

```{r generate_histogram_wage_growth_custom_bin, exercise=TRUE, exercise.cap="Practice: Generate a histogram of wage_growth with 50 bins"}
counties %>%
  filter(!is.na(wage_growth)) %>%
  ggplot(mapping = aes(x = wage_growth)) +
         geom_histogram(fill = "steelblue", color = "white", ___ = ___) +
         labs(title = "Wage Growth in US Counties in 2016", 
              x = NULL, 
              y = "Frequency", 
              caption = "Source: Linn, Nagler, Zilinsky")
```


```{r generate_histogram_wage_growth_custom_bin-hint-1}
# Hint 1 of 1
# The starter code has two blanks: ___ = ___.
# The first blank is the argument name: bins.
# The second blank is the value: 50.
# Write it as bins = 50 inside geom_histogram(),
# after color = "white" and separated by a comma.
```



```{r generate_histogram_wage_growth_custom_bin-check}
grade_this({
  code <- paste(.user_code, collapse = "\n")
  code_no_comments <- gsub("#.*", "", code)
  expr_text <- function(x) {
    if (is.null(x)) return("")
    paste(deparse(rlang::get_expr(x)), collapse = "")
  }

  if (!inherits(.result, "ggplot")) {
    fail("Your code should return a ggplot object. If you're seeing an error rather than a plot, the most common causes are: a missing + at the end of a line, unmatched parentheses, or a label string that is missing its quotation marks.")
  }
  if (!is.data.frame(.result$data) || !"wage_growth" %in% names(.result$data)) {
    fail("Use counties as the data frame in ggplot().")
  }
  x_expr <- expr_text(.result$mapping$x)
  if (!grepl("wage_growth", x_expr)) {
    fail("Map wage_growth to the x-axis.")
  }
  hist_layers <- Filter(function(layer) inherits(layer$stat, "StatBin"), .result$layers)
  if (length(hist_layers) == 0) {
    fail("Use geom_histogram() to create the histogram.")
  }
  has_50_bins <- any(vapply(hist_layers, function(layer) {
    identical(layer$stat_params$bins, 50)
  }, logical(1)))
  if (!has_50_bins && !grepl("bins\\s*=\\s*50", code_no_comments)) {
    fail("Set bins = 50 inside geom_histogram().")
  }
  has_white_outline <- any(vapply(hist_layers, function(layer) {
    identical(layer$aes_params$colour, "white") ||
      identical(layer$aes_params$color, "white")
  }, logical(1)))
  if (!has_white_outline) {
    fail('Set color = "white" inside geom_histogram() to outline the bars.')
  }

  pass("Good work! Notice how 50 bins reveals more detail than the default 30. Try the next exercise to see how binwidth works differently.")
})
```

Changing the number of bins changes how the data are grouped, which can noticeably alter the appearance of the histogram. With too few bins, the histogram can hide important features of the distribution because each bar covers a wide range of values. With too many bins, the plot can look noisy, making it difficult to see the overall pattern. Adjusting the number of bins helps you balance detail and readability so the histogram shows the meaningful structure in the data without overwhelming the viewer.


Alternatively, you can control the width of each bin directly. If a variable ranges from 0 to 100 and you set `binwidth = 10`, ggplot creates bins that each cover 10 units --- equivalent to setting `bins = 10` since 100/10 = 10. The two arguments give you two routes to the same result: `bins` sets how many bins to use, and `binwidth` sets how wide each bin is. Let's set the binwidth for wage growth to 0.8 --- eight-tenths of a percentage point.

**Explore: Generate the histogram for wage_growth with different binwidths**
```{r generate_histogram_wage_growth_custom_bin-width, exercise=TRUE,exercise.cap="Explore: Try different binwidths and see how the shape changes"}
counties %>%
  filter(!is.na(wage_growth)) %>%
  ggplot(mapping = aes(x = wage_growth)) +
         geom_histogram(fill = "steelblue", color = "white", binwidth = .8) +
         labs(title = "Wage Growth in US Counties in 2016", 
              x = NULL, 
              y = "Frequency", 
              caption = "Source: Linn, Nagler, Zilinsky")
```

Try experimenting with the binwidth in the plot to make sure you understand what this argument is doing.

With a binwidth of 0.8, the shape of the distribution is much clearer. Note something important: even with 50 bins or a narrow binwidth, **you cannot see the outlier counties** in this histogram. Their bars are so short compared to the central spike that they are invisible at this scale. You know they exist because of the range computed earlier --- this is one limitation of histograms. Rare extreme values can disappear visually even though they strongly affect statistics like the mean and range.

**Practice: Describe the shape of the distribution**

Describe the shape of the wage growth distribution based on the histogram you just produced.

```{r q-histogram-shape, echo=FALSE}
question("Based on the histogram of **wage_growth** with a narrow binwidth, which of the following best describes the shape of the distribution you can see in the plot?",
  correct = "Right. The distribution is roughly symmetric and bell-shaped.",
  answer("Roughly symmetric and bell-shaped, with most counties clustered near the center",
         correct = TRUE,
         message = "The distribution has a single peak near zero and tapers off fairly evenly on both sides. This is consistent with the mean and median being close to each other. The outlier counties identified by range() are not visible in this plot at this scale --- but the range you computed earlier tells you they are there. This is a key limitation of histograms: rare extreme values can disappear visually even when they strongly affect statistics like the mean and range."),
  answer("Asymmetric, with most counties near zero but a long tail of counties with very high wage growth",
         message = "The plot shows bars on both sides of the center that taper off fairly evenly in both directions. There is not a long tail extending only to one side --- the distribution is fairly balanced around the center."),
  answer("Uniform --- all values of wage growth are roughly equally common",
         message = "A uniform distribution would show bars of similar height across the full range. Here the bars are much taller near the center and shorter in the tails."),
  answer("Bimodal --- there are two distinct clusters of counties",
         message = "A bimodal distribution would show two separate peaks. The wage growth histogram has one central peak, not two."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  try_again = "Incorrect. Try again."
)
```

### Adding the median and mean to a histogram

Often it is useful to add a vertical line to a histogram to mark where the median or mean falls. The `geom_vline()` function adds this layer; its key argument is `xintercept`, which sets the x-axis value where the line appears. When the value is computed from a variable, you have two options: wrap the expression in `aes()` so bare variable names work (`aes(xintercept = median(wage_growth, na.rm = TRUE))`), or use the full `data$variable` syntax outside `aes()` (`xintercept = median(counties$wage_growth, na.rm = TRUE)`). Both produce the same result. Fixed settings like `color` and `linetype` always go outside `aes()`.

**Run and observe: Add the median line**

```{r try_it_label_median_line, exercise=TRUE,exercise.cap="Run and observe: Adding a median reference line with geom_vline()"}
counties %>%
  filter(!is.na(wage_growth)) %>%
  ggplot(mapping = aes(x = wage_growth)) +
         geom_histogram(fill = "steelblue", color = "white", bins = 60) +
         geom_vline(aes(xintercept = median(wage_growth, na.rm = TRUE)), color = "red", linetype = "dashed") +
         labs(title = "Wage Growth in US Counties in 2016", 
              x = NULL, 
              y = "Frequency", 
              caption = "Source: Linn, Nagler, Zilinsky")
```

The red dashed line marks the median of wage growth, sitting close to the center of the distribution. When the distribution is roughly symmetric, the median sits near the peak. In a right-skewed distribution, the median often lies to the right of the peak, while the mean is pulled farther toward the long right tail. Here, the median and mean are close to each other (4.70% and 4.33%), which is consistent with the roughly symmetric pattern visible in the histogram.

The line is useful but unlabeled. The `annotate()` function adds a text label at specific coordinates --- `x` and `y` place the center of the label on the plot axes, and `label` sets the text. Because `annotate()` does not inherit the data context from `ggplot()`, use `counties$wage_growth` rather than a bare variable name inside it. Use `round()` to limit decimal places, and `paste0()` to combine text with a computed value.

::: {.tip}
**Choosing annotation coordinates:** The `x` and `y` values in `annotate()` are on the same scale as the plot axes --- `x` is a value of the variable being plotted and `y` is a frequency count. To choose good coordinates: (1) run the plot first to see the axis ranges; (2) pick an `x` value near the line but offset so the text doesn't overlap it; (3) pick a `y` value above the tallest bars so the label is readable. If the annotation disappears or behaves strangely, your coordinates are probably outside the plot's visible range --- try values closer to the center. In your own analyses you will need to experiment to find coordinates that work. For this exercise, use `x = 25, y = 900`.
:::

Now add the mean line so you can see where the average falls relative to the rest of the distribution.

**Practice: Add a mean line and label to the histogram**

Fill in the one blank inside `geom_vline()` so the vertical line is placed at the mean value of wage growth. The annotation coordinates are provided. Two hints are available. Click "Submit Answer" to check your work.

```{r plot_vertical_and_text, exercise=TRUE, exercise.cap="Practice: Add a mean line and annotation to the histogram"}
counties %>%
  filter(!is.na(wage_growth)) %>%
  ggplot(mapping = aes(x = wage_growth)) +
         geom_histogram(fill = "steelblue", color = "white", bins = 60) +
         geom_vline(aes(xintercept = ___(wage_growth, na.rm = TRUE)), color = "red", linetype = "dotted") +
         annotate("text", x = 25, y = 900, 
                  label = paste0("Mean = ", 
                                 round(___(counties$wage_growth, na.rm = TRUE), 3))) +
         labs(title = "Wage Growth in US Counties in 2016", 
              x = NULL, 
              y = "Frequency", 
              caption = "Source: Linn, Nagler, Zilinsky")
```

```{r plot_vertical_and_text-hint-1}
# Hint 1 of 2
# Inside geom_vline(), you need a function that
# computes the average of wage_growth.
# The same function you used earlier for central
# tendency, with na.rm = TRUE.
```

```{r plot_vertical_and_text-hint-2}
# Hint 2 of 2
# Inside geom_vline(), the xintercept aesthetic should
# be set to the function that computes the arithmetic
# average of wage_growth, with na.rm = TRUE.
# This is the same function you used in the earlier
# central tendency exercise.
```


```{r plot_vertical_and_text-check}
grade_this({
  code <- paste(.user_code, collapse = "\n")
  code_no_comments <- gsub("#.*", "", code)
  expr_text <- function(x) {
    if (is.null(x)) return("")
    paste(deparse(rlang::get_expr(x)), collapse = "")
  }

  if (!inherits(.result, "ggplot")) {
    fail("Your code should return a ggplot object. If you're seeing an error rather than a plot, the most common causes are: a missing + at the end of a line, unmatched parentheses, or a label string that is missing its quotation marks.")
  }
  vline_layers <- Filter(function(layer) inherits(layer$geom, "GeomVline"), .result$layers)
  if (length(vline_layers) == 0) {
    fail("Add a geom_vline() layer to draw the mean line.")
  }
  expected_mean <- mean(counties$wage_growth, na.rm = TRUE)
  has_mean_line <- any(vapply(vline_layers, function(layer) {
    x_expr <- expr_text(layer$mapping$xintercept)
    fixed_x <- layer$geom_params$xintercept
    grepl("mean\\s*\\(", x_expr) && grepl("wage_growth", x_expr) ||
      (!is.null(fixed_x) && isTRUE(all.equal(as.numeric(fixed_x), expected_mean)))
  }, logical(1)))
  if (!has_mean_line) {
    fail("Set xintercept to mean(wage_growth, na.rm = TRUE) inside geom_vline().")
  }
  has_dotted_line <- any(vapply(vline_layers, function(layer) {
    identical(layer$aes_params$linetype, "dotted")
  }, logical(1)))
  if (!has_dotted_line) {
    fail('Set linetype = "dotted" for the mean line.')
  }
  has_red_line <- any(vapply(vline_layers, function(layer) {
    identical(layer$aes_params$colour, "red") ||
      identical(layer$aes_params$color, "red")
  }, logical(1)))
  if (!has_red_line) {
    fail('Set color = "red" for the mean line.')
  }
  text_layers <- Filter(function(layer) inherits(layer$geom, "GeomText"), .result$layers)
  if (length(text_layers) == 0 && !grepl("annotate\\s*\\(", code_no_comments)) {
    fail("Add an annotate() layer to label the mean line.")
  }
  has_coords <- grepl("x\\s*=\\s*25", code_no_comments) && grepl("y\\s*=\\s*900", code_no_comments)
  if (!has_coords) {
    fail("Use x = 25 and y = 900 in annotate().")
  }
  if (!grepl("Mean", code_no_comments)) {
    fail('Use a label that identifies the line as the "Mean".')
  }

  pass("Well done! The mean line lands near the center of the distribution, consistent with the roughly symmetric shape we described. In your own work, you will need to try different x and y coordinates until the label sits clearly above the bars without overlapping the line.")
})
```


Now return to **dem2p_percent** and use a histogram to check the clue you got from the mean and median. If the mean is higher than the median, you should expect a right-skewed distribution.

**Practice: Histogram for dem2p_percent**

This exercise combines four layers. The first two --- `ggplot()` and `geom_histogram()` --- you have used before. The last two --- `geom_vline()` and `annotate()` --- you ran in the previous exercises but are writing from scratch for the first time here. Build incrementally: get the histogram working first, then add the vertical line and run again, then add the annotation.

Build a histogram of **dem2p_percent** with a median reference line and annotation. Filter out missing values before plotting. Four hints are available. Click "Submit Answer" to check your work.

Your plot should have:

- Number of bins: 30
- Histogram bar outlines: `color = "blue"`
- A vertical median line: red, dotted
- An annotation at `x = 55, y = 280` labeling the line with `"Median = "` and the rounded median value using `paste0()`
- Title: "Democratic Vote Share"
- No subtitle --- unlike `x` and `y` axis labels, which default to the mapped variable name and must be set to `NULL` to suppress them, `subtitle` defaults to nothing and can simply be omitted from `labs()`
- x-axis label: "Percent of Two-Party Vote"
- y-axis label: "Frequency"
- Caption: "Source: Linn, Nagler, Zilinsky"

```{r histogram_dem2p_percent, exercise=TRUE, exercise.cap="Practice: Histogram of dem2p_percent with a median line"}

```

```{r histogram_dem2p_percent-hint-1}
# Hint 1 of 4
# Before passing data to ggplot(), filter out
# missing values --- you have done this in every
# histogram and proportion bar chart so far.
# Then map dem2p_percent to the x-axis and add
# geom_histogram(). Run the code to confirm you
# have a working histogram before adding anything else.
```

```{r histogram_dem2p_percent-hint-2}
# Hint 2 of 4
# Add a geom_vline() layer for the median line.
# Two approaches work:
# Option 1: geom_vline(aes(xintercept = median(dem2p_percent, na.rm = TRUE)))
# Option 2: geom_vline(xintercept = median(counties$dem2p_percent, na.rm = TRUE))
# Set color = "red" and linetype = "dotted" outside aes().
```

```{r histogram_dem2p_percent-hint-3}
# Hint 3 of 4
# Add annotate("text", x = 55, y = 280,
#              label = paste0("Median = ",
#                             round(median(counties$dem2p_percent,
#                                         na.rm = TRUE), 3)))
# Use counties$dem2p_percent inside annotate() ---
# it does not inherit the data context from ggplot().
```

```{r histogram_dem2p_percent-hint-4}
# Hint 4 of 4
# Full pipeline:
# counties %>%
#   filter(!is.na(dem2p_percent)) %>%
#   ggplot(mapping = aes(x = dem2p_percent)) +
#   geom_histogram(color = "blue", bins = 30) +
#   geom_vline() with the median line +
#   annotate("text", x = 55, y = 280,
#            label = paste0("Median = ", ...)) +
#   labs(title = "Democratic Vote Share",
#        x = "Percent of Two-Party Vote",
#        y = "Frequency",
#        caption = "Source: Linn, Nagler, Zilinsky")
```


```{r histogram_dem2p_percent-check}
grade_this({
  code <- paste(.user_code, collapse = "\n")
  code_no_comments <- gsub("#.*", "", code)
  expr_text <- function(x) {
    if (is.null(x)) return("")
    paste(deparse(rlang::get_expr(x)), collapse = "")
  }

  if (!inherits(.result, "ggplot")) {
    fail("Your code should return a ggplot object. If you're seeing an error rather than a plot, the most common causes are: a missing + at the end of a line, unmatched parentheses, or a label string that is missing its quotation marks.")
  }
  if (any(is.na(.result$data$dem2p_percent))) {
    fail("Filter out missing values with filter(!is.na(dem2p_percent)) before plotting.")
  }
  if (!is.data.frame(.result$data) || !"dem2p_percent" %in% names(.result$data)) {
    fail("Use dem2p_percent from the counties data frame for this plot.")
  }

  expected_values <- sort(as.numeric(counties$dem2p_percent[!is.na(counties$dem2p_percent)]))
  actual_values   <- sort(as.numeric(.result$data$dem2p_percent))
  if (!isTRUE(all.equal(actual_values, expected_values, tolerance = 1e-8))) {
    fail("Use all nonmissing observations of dem2p_percent from the counties data frame. Do not remove any valid observations.")
  }
  x_expr <- expr_text(.result$mapping$x)
  if (!grepl("dem2p_percent", x_expr)) {
    fail("Map dem2p_percent to the x-axis.")
  }
  hist_layers <- Filter(function(layer) inherits(layer$stat, "StatBin"), .result$layers)
  if (length(hist_layers) == 0) {
    fail("Use geom_histogram() to create the histogram.")
  }
  has_blue_outline <- any(vapply(hist_layers, function(layer) {
    identical(layer$aes_params$colour, "blue") ||
      identical(layer$aes_params$color, "blue")
  }, logical(1)))
  if (!has_blue_outline) {
    fail('Set color = "blue" inside geom_histogram() to outline the bars.')
  }
  vline_layers <- Filter(function(layer) inherits(layer$geom, "GeomVline"), .result$layers)
  if (length(vline_layers) == 0) {
    fail("Add a geom_vline() layer for the median line.")
  }
  expected_median <- median(counties$dem2p_percent, na.rm = TRUE)
  has_median_line <- any(vapply(vline_layers, function(layer) {
    xint_expr <- expr_text(layer$mapping$xintercept)
    fixed_x <- tryCatch(as.numeric(layer$geom_params$xintercept[1]), error = function(e) NA)
    in_aes <- grepl("median\\s*\\(", xint_expr) && grepl("dem2p_percent", xint_expr)
    in_code <- grepl("median\\s*\\(.*dem2p_percent", code_no_comments)
    in_fixed <- !is.na(fixed_x) && isTRUE(all.equal(round(fixed_x, 2), round(expected_median, 2)))
    in_aes || in_code || in_fixed
  }, logical(1)))
  if (!has_median_line) {
    fail("Set xintercept to the median of dem2p_percent inside geom_vline(). You can use aes(xintercept = median(dem2p_percent, na.rm = TRUE)) or xintercept = median(counties$dem2p_percent, na.rm = TRUE).")
  }
  has_dotted_line <- any(vapply(vline_layers, function(layer) {
    identical(layer$aes_params$linetype, "dotted")
  }, logical(1)))
  if (!has_dotted_line) {
    fail('Set linetype = "dotted" for the median line.')
  }
  has_red_line <- any(vapply(vline_layers, function(layer) {
    identical(layer$aes_params$colour, "red") ||
      identical(layer$aes_params$color, "red")
  }, logical(1)))
  if (!has_red_line) {
    fail('Set color = "red" for the median line.')
  }
  text_layers <- Filter(function(layer) inherits(layer$geom, "GeomText"), .result$layers)
  if (length(text_layers) == 0 && !grepl("annotate\\s*\\(", code_no_comments)) {
    fail("Add an annotate() layer to label the median line.")
  }
  has_coords <- grepl("x\\s*=\\s*55", code_no_comments) && grepl("y\\s*=\\s*280", code_no_comments)
  if (!has_coords) {
    fail("Use x = 55 and y = 280 in annotate().")
  }
  if (!grepl("Median", code_no_comments)) {
    fail('Use a label that identifies the line as the median, such as "Median = " followed by the rounded value using paste0().')
  }

  if (!identical(.result$labels$title, "Democratic Vote Share")) {
    fail('Check the plot title. Use "Democratic Vote Share".')
  }
  if (!identical(.result$labels$x, "Percent of Two-Party Vote")) {
    fail('Set the x-axis label to "Percent of Two-Party Vote".')
  }
  if (!identical(.result$labels$y, "Frequency")) {
    fail('Set the y-axis label to "Frequency".')
  }
  if (!identical(.result$labels$caption, "Source: Linn, Nagler, Zilinsky")) {
    fail('Check the caption. Use "Source: Linn, Nagler, Zilinsky".')
  }

  pass("Well done! The histogram reveals the shape of the distribution. Compare it to the box plot you will create below --- both show the skew, but in different ways.")
})
```

The distribution is clearly **right-skewed**. A large concentration of counties have Democratic vote percentages well below the median, while fewer counties gave high percentages of their two-party vote to the Democratic candidate. The long right tail shows that only a small subset of counties cast very large percentages of their two-party votes for Clinton.

### The box plot

A [**box plot**]{.important-text} is another way to visualize the distribution of an interval variable. It is especially useful when you want to compare the distribution of one variable across categories of another variable. A box plot brings the median, the IQR, and outlying observations together in a compact visual summary.

To generate a box plot, change the basic histogram plot in two ways. First, map the variable to the y-axis rather than the x-axis. This is because a box plot displays the distribution vertically --- the median line runs horizontally across the box, the whiskers extend up and down, and outliers appear as dots above or below. If the variable were mapped to x instead, the plot would be rotated 90 degrees and harder to read in the standard orientation. Use `x = ""` to give the box a standard width and suppress the unnecessary x-axis label --- without it, the box stretches across the full plot area. Second, use `geom_boxplot()` rather than `geom_histogram()`.

Let's plot the percent of working-age adults. Recall this variable is named **employ_pop_ratio_25_64**. 

**Run and observe: Plot the employment-to-population ratio with a box plot**
```{r boxplot_employment_ratio, exercise=TRUE,exercise.cap="Run and observe: A box plot of the employment-to-population ratio"}
counties %>%
  filter(!is.na(employ_pop_ratio_25_64)) %>%
  ggplot(aes(x = "", y = employ_pop_ratio_25_64)) +
  geom_boxplot() +
  labs(title = "Distribution of the Employment to Population Ratio \n Adults Aged 25-64", 
       subtitle = "US Counties, 2016",
       x = NULL, 
       y = "Percent",
       caption = "Source: Linn, Nagler, Zilinsky") 
```
1. The heavy horizontal line inside the box indicates the median --- the value at the 50th percentile. Here the median employment-to-population ratio is about 66.8%.
2. The boxed area shows the values in the middle 50% of the data --- from the 25th to the 75th percentile of the distribution. The height of the box (the value of the 3rd quartile minus the value of the 1st quartile) is the interquartile range or IQR. It provides a measure of the spread of the middle 50% of values. Here, the middle half of counties have employment-to-population ratios between about 51% and 84%. Longer boxes would indicate more variability.
3. The “whisker” portions of these plots extend out from the bottoms and tops of the boxes and represent points less than the 25th percentile and greater than the 75th percentiles, respectively. They’re set to extend out no more than 1.5 × IQR units away from either end of the boxes. If there are no observations that far out, the whiskers end at the nearest extreme data point. Longer whiskers indicate more variability. 
4. Observations outside that range -- beyond 1.5 × IQR from the box -- are indicated by circles. The circles represent outliers or unusual values. Here you see that several counties have employment-to-population ratios above 100%. This can happen when a county has more workers employed there than working-age residents living there --- in other words, when the county draws workers who commute in from elsewhere.


**Practice: Create a box plot for dem2p_percent**

Produce the box plot, filtering out missing values before plotting. Then look at where the median line sits within the box and whether the whiskers extend evenly before reading the interpretation below. Two hints are available. Click "Submit Answer" to check your work.

Your plot should have:

- Title: "Democratic Vote Share"
- x-axis label: omitted (`NULL`)
- y-axis label: "Percent of Two-Party Vote"
- Caption: "Source: Linn, Nagler, Zilinsky"

```{r boxplot_dem2p, exercise=TRUE, exercise.cap="Practice: Create a box plot for dem2p_percent"}

```

```{r boxplot_dem2p-hint-1}
# Hint 1 of 2
# Use ggplot() with aes(x = "", y = dem2p_percent)
# from the counties data frame.
# Add geom_boxplot() and labs() with the title
# "Democratic Vote Share" and y = "Percent of Two-Party Vote".
```

```{r boxplot_dem2p-hint-2}
# Hint 2 of 2
# Map the variable to y, not x --- box plots for a
# single variable use y so the distribution runs
# vertically with the median line horizontal.
# Use x = "" to keep the plot clean without an
# unnecessary x-axis category label.
```


```{r boxplot_dem2p-check}
grade_this({
  code <- paste(.user_code, collapse = "\n")
  code_no_comments <- gsub("#.*", "", code)
  expr_text <- function(x) {
    if (is.null(x)) return("")
    paste(deparse(rlang::get_expr(x)), collapse = "")
  }

  if (!inherits(.result, "ggplot")) {
    fail("Your code should return a ggplot object. If you're seeing an error rather than a plot, the most common causes are: a missing + at the end of a line, unmatched parentheses, or a label string that is missing its quotation marks.")
  }
  if (!is.data.frame(.result$data) || !"dem2p_percent" %in% names(.result$data)) {
    fail("Use dem2p_percent from the counties data frame for this plot.")
  }
  if (any(is.na(.result$data$dem2p_percent))) {
    fail("Filter out missing values with filter(!is.na(dem2p_percent)) before plotting.")
  }

  expected_values <- sort(as.numeric(counties$dem2p_percent[!is.na(counties$dem2p_percent)]))
  actual_values   <- sort(as.numeric(.result$data$dem2p_percent))
  if (!isTRUE(all.equal(actual_values, expected_values, tolerance = 1e-8))) {
    fail("Use all nonmissing observations of dem2p_percent from the counties data frame. Do not remove any valid observations.")
  }
  x_expr <- expr_text(.result$mapping$x)
  if (!identical(x_expr, "\"\"") && !identical(x_expr, "''")) {
    fail('Use x = "" inside aes() for a cleaner one-variable box plot.')
  }
  y_expr <- expr_text(.result$mapping$y)
  if (!grepl("dem2p_percent", y_expr)) {
    fail("Map dem2p_percent to the y-axis.")
  }
  has_boxplot <- any(vapply(.result$layers, function(layer) {
    inherits(layer$geom, "GeomBoxplot")
  }, logical(1)))
  if (!has_boxplot) {
    fail("Use geom_boxplot() to create the box plot.")
  }
  if (!identical(.result$labels$title, "Democratic Vote Share")) {
    fail('Use "Democratic Vote Share" as the title.')
  }
  if (!identical(.result$labels$y, "Percent of Two-Party Vote")) {
    fail('Set the y-axis label to "Percent of Two-Party Vote".')
  }
  if (!identical(.result$labels$caption, "Source: Linn, Nagler, Zilinsky")) {
    fail('Check the caption. Use "Source: Linn, Nagler, Zilinsky".')
  }

  pass("Good work! The chart is correct.")
})
```

The box plot for Democratic vote percentage shows a median around 30%, consistent with what you computed earlier. The IQR box spans roughly 20% to 45%, meaning the middle half of counties gave Clinton between 20% and 45% of the two-party vote. The long upper whisker and outlier circles above it indicate that a small number of counties gave Clinton very high vote percentages --- consistent with the right-skewed distribution you identified from the mean exceeding the median.

## Check Your Understanding

Now let’s step back and practice choosing among the tools from this tutorial. The questions below ask you to connect the level of measurement to the appropriate summary statistic or plot, and then to interpret what those summaries tell you about the center, spread, and shape of a distribution.

**Apply: Identify core distribution concepts**

Think carefully about each statement before selecting.

```{r identify_distribution_concepts, echo=FALSE}
question("Which of the following statements are true? Select all that apply.",
  correct = "Right. Both statements are true.",
  answer("The mean is a robust summary statistic",
         message = "The mean is not robust --- it is sensitive to outliers. A single extreme value can pull the mean far from the center of the distribution. The median is the robust alternative."),
  answer("A distribution whose shape is skewed right has a long tail on the right", correct = TRUE,
         message = "The direction of skew is named for the tail, not the peak. A right-skewed distribution has most values clustered on the left and a long tail extending to the right. The terminology can be counterintuitive because the peak is on the left, not the right --- it is the tail that gives skew its direction."),
  answer("The standard deviation measures the typical distance of observations from the mean", correct = TRUE,
         message = "The standard deviation is the typical distance of observations from the mean in the original units of the variable. A larger standard deviation means observations are more dispersed; a smaller one means they cluster more tightly around the mean."),
  answer("The mode should be reported as a measure of central tendency regardless of a variable's level of measurement",
         message = "The mode is generally not useful for interval variables, where values are so precisely measured that no single value repeats often enough to be informative. For interval data, the median and mean are the appropriate measures of central tendency."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  try_again = "Incorrect. Try again."
)
```

Consider the following histogram and answer the question below:

```{r, echo=FALSE, fig.alt="Histogram showing public expenditure on education as a percentage of GDP for countries worldwide. The distribution is right-skewed with a peak near 4 to 5 percent of GDP, most countries spending between 2 and 7 percent, and two isolated outliers near 12 percent. A text annotation shows Standard Deviation equals 1.92."}
world %>%
  filter(!is.na(spendeduc)) %>%
ggplot(aes(x = spendeduc)) +
  geom_histogram(fill = "steelblue", color = "white", bins = 20) +
  labs(title = "Public Expenditure on Education Around the World",
       caption = "Country Data from the United Nations",
       x = "Percent of GDP",
       y = "Frequency") +
          annotate("text", x = 10, y = 18, 
                  label = paste0("Standard Deviation = ", 
                                 round(sd(world$spendeduc, na.rm = TRUE), 2))) 
```

**Apply: Interpret the education-spending histogram**

Based on what you learned about reading histograms, interpret the plot above and reason about what the shape implies for which statistics are most appropriate.

```{r plot_interpretation, echo=FALSE}
question("Based on the histogram of public education spending, which of the following are true? Select all that apply.",
  correct = "Right. Both statements are true.",
  answer("The outliers near 12 percent of GDP are likely pulling the mean above the median", correct = TRUE,
         message = "Extreme values in the upper tail can pull the mean upward much more strongly than they affect the median. That is the defining consequence of right skew with high outliers: the mean is pulled away from the center while the median stays much closer to where most observations sit."),
  answer("The median is a more robust summary of the typical country's education spending than the mean", correct = TRUE,
         message = "The median is much less affected by the countries spending 10 percent or more, while the mean is pulled upward by them. When outliers are present and the distribution is skewed, the median better represents what is typical."),
  answer("The mean is likely a better measure of central tendency than the median",
         message = "The right tail contains a small number of countries spending much more than typical. Those values pull the mean upward much more strongly than they affect the median, making the median the more robust measure here."),
  answer("The distribution appears bimodal, with two distinct clusters of countries at different spending levels",
         message = "The histogram shows a single peak near 4 to 5 percent of GDP. The isolated bars near 12 percent are outliers, not a second peak of comparable size."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  try_again = "Incorrect. Try again."
)
```

## The Takeaways

Interval variables give you a wider set of descriptive tools because their values are ordered and the distances between values are meaningful. With Tutorial 4 covering nominal variables, Tutorial 5 covering ordinal variables, and this tutorial completing the sequence, you now have a full toolkit for describing one variable at any level of measurement. The same three dimensions apply throughout --- central tendency, dispersion, and shape --- but the tools available depend on what the values actually mean. The median requires a meaningful rank order, so it is not appropriate for nominal variables. The mean assumes equal intervals between values, so it is not appropriate for ordinal variables. Only at the interval level are all three measures of central tendency available.

- [**Central tendency**]{.important-text} for interval variables is described by the [**median**]{.important-text} and [**mean**]{.important-text}, computed with `median()` and `mean()`, usually with `na.rm = TRUE`. The mode is almost never useful for interval data because values are often measured so precisely that no single value repeats very often. The median is the value at the 50th percentile and is robust to outliers. The mean is the arithmetic average and can be pulled away from the center by extreme values. When the distribution is roughly symmetric, the two are similar. When they diverge, the median is usually the better summary of what is typical.

- [**Dispersion**]{.important-text} for interval variables is captured by the [**range**]{.important-text} (`range()`), [**percentiles**]{.important-text} (`quantile()`), and [**standard deviation**]{.important-text} (`sd()`). The interquartile range is the distance from the 25th to the 75th percentile; it describes the middle half of the data and is less affected by outliers than the full range. The standard deviation is the typical distance of observations from the mean, in the original units of the variable. One important caution: standard deviations are only comparable within a variable or across samples of the same variable.

- [**Shape**]{.important-text} for interval variables is best assessed visually. A [**histogram**]{.important-text} (`geom_histogram()`) groups values into bins and shows the overall pattern of the distribution --- its center, spread, symmetry or skew, and whether it is unimodal, bimodal, or uniform. The `bins` and `binwidth` arguments affect how much detail you see. A [**box plot**]{.important-text} (`geom_boxplot()`) summarizes the median, IQR, and outliers in a compact form that becomes especially useful when comparing distributions across groups. When reading either plot, ask: where is the center, how spread out are the values, is the distribution symmetric or skewed, and are there outliers?

In R, load [dplyr]{.package-name} and [ggplot2]{.package-name} once per session with `library(dplyr)` and `library(ggplot2)` at the top of your file 

For visualization, histograms and box plots follow the same layered ggplot structure as bar charts. Always filter out missing values before passing data to `ggplot()` --- if you do not, ggplot will print a warning message that would appear in a knitted document.

- **Histograms:** Use `geom_histogram()` in place of `geom_bar()`, with `fill` to set the bar color, `color = "white"` to add bar outlines, and `bins` or `binwidth` to control bin size.

- **Reference lines:** Use `geom_vline()` to mark a statistic. Inside `aes()`, bare variable names work: `aes(xintercept = median(wage_growth, na.rm = TRUE))`. Outside `aes()`, use the full `data$variable` syntax. Fixed settings like `color` and `linetype` always go outside `aes()`. To label the line, add `annotate("text", x = ..., y = ..., label = ...)` --- use `counties$variable` inside `annotate()` since it does not inherit the data context from `ggplot()`.

- **Box plots:** Map the variable to `y` rather than `x`, use `geom_boxplot()`, and set `x = ""` to suppress the unnecessary x-axis category label.

Which summaries should you present? Report the mean or median and the standard deviation or IQR, or show a histogram or box plot --- but do not bury your reader in every statistic and every graph just because you can. Tables and plots often communicate overlapping information, so choose the display that best serves your purpose rather than presenting both automatically. Whatever you choose, describe its important features in the text --- do not leave a table or figure to speak for itself. The goal is always to give readers enough to see the center, spread, and shape of the distribution.

The figure below brings together the tools for describing [**nominal**]{.important-text}, [**ordinal**]{.important-text}, and [**interval variables**]{.important-text} across all three tutorials. Use it as a reference as you move into bivariate analysis.

```{r univariate-description-figure, echo=FALSE,fig.align = "center", out.width = "95%",  fig.alt = "Infographic titled 'Univariate Description'. Left sidebar lists Central Tendency, Dispersion, Distribution Shape, and Visualization. Three main panels show stacked blue boxes for interval, nominal, and ordinal level variables with their recommended measures of central tendency, dispersion, distribution shape, and visualization."}
knitr::include_graphics("images/univariate_description-4.png")
```

## Submit Your Work

When you have completed all exercises in this tutorial, generate your submission code using the button below. Copy the full code that appears and paste it into the Canvas submission box for this tutorial. The code can be long --- make sure you have copied all of it before submitting.

```{r encode-logic-06, context="server"}
learnrhash::encoder_logic()
```

```{r encode-ui-06, echo=FALSE}
learnrhash::encoder_ui(
  ui_before = shiny::div(
    shiny::p("Click the button below to generate your submission code.
              Copy the full code and paste it into Canvas.")
  )
)
```

```{r tutorial-success-banner-06, results="asis", echo=FALSE}
cat('
<p class="tutorial-complete"
   role="status"
   aria-label="Well done! You have completed this tutorial and can now describe and visualize interval variables using measures of central tendency, dispersion, and distribution shape.">
  <span aria-hidden="true">📊</span>
  Well done! You can now describe and visualize interval variables using measures of central tendency, dispersion, and distribution shape.
</p>
')
```
