---
title: "Hypothesis Tests: Two Categorical Variables"
tutorial:
  id: "12-Hypotheses-Categorical"
output: 
  QPATutorialsCourse::qpa_tutorial:
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
description: >
  Learn how to conduct a hypothesis test of the relationship between two categorical variables.
---

## Overview {#main-content}

Tutorials 7, 8, and 9 described bivariate relationships --- cross-tabs, group means, scatter plots --- but every example ended with the same caution: a descriptive pattern consistent with a hypothesis is not the same as evidence that the pattern is real. In Tutorials 10 and 11, you built the tools a formal test requires: the sampling distribution, the p-value, and the decision rule for rejecting or failing to reject a null hypothesis. Tutorials 12, 13, and 14 put those tools to work. This tutorial takes the cross-tab relationships from Tutorial 7 and subjects them to formal **hypothesis tests** --- using sample data to assess whether an observed relationship would be unlikely if there were truly no relationship in the population. Tutorial 13 does the same for the group-mean relationships from Tutorial 8, and Tutorial 14 for the correlations from Tutorial 9.

In this tutorial, you will learn how to test whether two categorical variables are related in the population. Specifically, we will cover:

* How to test whether two categorical variables are related using a chi-square test of independence, with the `CrossTable()` function in the [descr]{.package-name} package
* How to sharpen a general chi-square result into a directional difference-of-proportions test using `prop.test()`, when a hypothesis predicts which of two specific groups will have the higher proportion
* How to organize every hypothesis test into five explicit steps, from stating the null hypothesis through deciding what the result means substantively
* How to test a different kind of question --- whether a single sample proportion differs from a fixed hypothesized value --- using the same `prop.test()` function with no second group involved


```{r setup, include=FALSE}
library(learnr)
library(descr)
library(gradethis)
tutorial_options(exercise.checker = gradethis::grade_learnr,
                 exercise.completion = FALSE)
knitr::opts_chunk$set(error = TRUE)
options(lifecycle_verbosity = "quiet")
options(cli.num_colors = 1)
data("counties", package = "QPATutorialsCourse")
data("world", package = "QPATutorialsCourse")
```


## Testing the Relationship Between Two Categorical Variables

In Tutorial 7, you read a cross-tab by comparing column proportions across categories of an explanatory variable. If those proportions differed in the direction a hypothesis predicted, the pattern was consistent with the hypothesis. You found that Trump won a majority in a steadily larger share of counties as counties grew more rural, that majority-white counties gave him a majority far more often than counties with any other racial majority did, and that dictatorships were less likely than democracies to have a high degree of gender equality. Each of those patterns pointed where its hypothesis predicted. What description alone cannot tell us is whether differences of that size are larger than chance alone would produce.

A cross-tab cannot answer that. A sample is just one sample --- if we drew a different sample from the same population, the proportions would differ at least a little, just by chance. The question is whether a gap this large would be surprising if the two variables were truly unrelated in the population, and answering it is what a hypothesis test is for.

A word about the data frames in this tutorial. Neither **counties** nor **world** is a random sample from a larger population the way a survey is. As Tutorial 11 explained, inference is still appropriate when the research question is about a relationship --- a systematic pattern that could have looked different under different conditions. A p-value answers: if there were truly no relationship, how surprising would a pattern this strong be in these data?

We test whether two categorical variables are related using two complementary tools. The **chi-square test of independence** considers every category at once and asks whether the variables are related at all. A **difference-of-proportions test** asks a narrower but sharper question: comparing exactly two specific groups, it can test a specific directional prediction in a way the chi-square test cannot. Since building a cross-tab already gives us everything the chi-square test needs, we test the general claim first, then sharpen to a directional two-group test when the hypothesis supports one. We also look at a related but different question: whether a single sample proportion differs from some fixed hypothesized value.

## The Chi-Square Test of Independence

The [**chi-square test of independence**]{.important-text} starts with a question: what would the cross-tab look like if the two variables had no relationship at all --- if they were completely independent? We can calculate exactly that. For any cell in the table, the [**expected count**]{.important-text} under [**independence**]{.important-text} is:

$$E=\frac{(\text{row total})(\text{column total})}{n}$$

where $n$ is the total number of cases in the table. This formula says: if the two variables are independent, each cell should contain its fair share of cases, in proportion to the row and column totals.

Let's see this with a small example. Suppose 100 people were each categorized by two variables: Group (A or B) and Outcome (Yes or No). Fifty people are in Group A and fifty are in Group B; 60 people said Yes and 40 said No, overall. If Group and Outcome were independent, here is what we would expect to see:

[Expected Counts Under Independence]{.section-header}

| | Yes | No | Total |
|:--|:--:|:--:|:--:|
| **Group A** | 30 | 20 | 50 |
| **Group B** | 30 | 20 | 50 |
| **Total** | 60 | 40 | 100 |

Now suppose this is what we actually observed:

[Observed Counts]{.section-header}

| | Yes | No | Total |
|:--|:--:|:--:|:--:|
| **Group A** | 40 | 10 | 50 |
| **Group B** | 20 | 30 | 50 |
| **Total** | 60 | 40 | 100 |

Group A said "Yes" more often than independence would predict (40 observed vs. 30 expected), and Group B said "Yes" less often than independence would predict (20 observed vs. 30 expected). This is exactly the kind of pattern a chi-square test is built to detect.

### Calculating the Chi-Square Statistic

The [**chi-square statistic**]{.important-text}, denoted $\chi^2$, summarizes how far the observed counts are from the expected counts, across every cell in the table:

$$\chi^2=\Sigma \frac{(O-E)^2}{E}$$

where $O$ is the observed count in a cell and $E$ is the expected count in that same cell, calculated as above.

Think through what this formula tells us. For each cell, we square the gap between what we observed and what independence would predict, then divide by the expected count, so that cells with naturally larger expected counts do not dominate just because they are bigger. We add that result up across every cell in the table. When observed counts are close to expected counts everywhere, each term is small, and $\chi^2$ is small. When observed counts differ substantially from expected counts, $\chi^2$ is large.

For example, the Group A/Yes cell contributes $(40-30)^2/30=3.33$ to the total, and the Group B/Yes cell contributes $(20-30)^2/30=3.33$. The two No cells are off by the same 10, but each is measured against an expected count of 20 rather than 30: Group A/No contributes $(10-20)^2/20=5.00$ and Group B/No contributes $(30-20)^2/20=5.00$. The same gap counts for more when what independence predicted was smaller. Adding all four gives $\chi^2=16.67$ for this table.

In Tutorials 10 and 11, the sampling distribution you built intuition for was the normal curve --- symmetric, bell-shaped, centered on the population parameter. $\chi^2$ follows a different distribution entirely. Because it is built from squared terms, $\chi^2$ can never be negative, so its distribution starts at 0 and stretches out in a long tail to the right, rather than spreading symmetrically in both directions the way the normal curve does. The exact shape depends on the table's degrees of freedom, which is why R reports them alongside the test statistic --- but the precise calculation is one more thing we will not dwell on in this course.

What does not change is the underlying logic. Just as with the normal distribution, a hypothesis test locates your observed statistic on this distribution and asks how far into the tail it falls. A $\chi^2$ value far out in that long right tail is one that would be rare if the two variables were truly independent --- exactly the kind of evidence a hypothesis test is built to evaluate. R converts $\chi^2$ into a p-value this way automatically, and we apply the same decision rule as always: a small p-value (conventionally below .05) leads us to reject the null hypothesis of independence, while a large p-value means we fail to reject it.

### Running the Test with `CrossTable()`

We already know how to build a cross-tab with the `CrossTable()` function in the [descr]{.package-name} package. The `x` argument takes the **outcome variable** and `y` takes the **explanatory variable** --- `x` defines the rows and `y` defines the columns, so with `prop.c = TRUE` each column proportion shows how the outcome distribution varies across categories of the explanatory variable. To add the chi-square test, we set the argument `chisq = TRUE`. This adds the chi-square statistic, its degrees of freedom, and its p-value to the output; the statistic and the p-value are the two we use to reach a conclusion.

We will also set `expected = TRUE`, which adds the expected count --- calculated with the formula above --- to every cell.

Putting those together, every chi-square cross-tab in this tutorial has the same shape. All eight are worth writing out each time: left to its defaults, `CrossTable()` prints five numbers in every cell and runs no test at all, so every setting after `x` and `y` is doing work.

**Example code:**

```{r crosstable-template, exercise=FALSE, eval=FALSE}
CrossTable(x = yourdata$outcome_variable,
           y = yourdata$explanatory_variable,
           prop.r = FALSE,
           prop.c = TRUE,
           prop.t = FALSE,
           prop.chisq = FALSE,
           chisq = TRUE,
           expected = TRUE)
```

One thing about this test is worth noting before we formalize the procedure. Its hypotheses do not name a number. The null is that the two variables are independent in the population and the alternative is that they are related, where most tests state a null about a specific parameter value. The logic is unchanged --- both are still claims about the population rather than about the sample --- but the chi-square test is the one place in this tutorial where you will not see a parameter in $H_0$.

## Applying the Five-Step Hypothesis Testing Procedure

In Tutorial 11, you learned the building blocks of a hypothesis test: a null and alternative hypothesis, a test statistic, a p-value, and the decision rule for rejecting or failing to reject the null hypothesis. Here, we organize those same building blocks into a [**five-step procedure**]{.important-text}, which we will follow for every hypothesis test in this tutorial and in those that follow:

1. State the null hypothesis, $H_0$ --- a claim about a population parameter.

2. State the alternative hypothesis, $H_A$. Like the null, it is a claim about the population.

3. Run the test and obtain a test statistic.

4. Find the p-value associated with that test statistic.

5. Decide whether to reject or fail to reject $H_0$, and report what that means substantively.

::: {.tip}
**Hypotheses are always about the population.** A hypothesis test asks whether a pattern in your sample reflects something real in the population. For this reason, $H_0$ and $H_A$ are always stated in terms of population parameters rather than sample statistics --- written in Greek letters, alongside the Roman-letter sample statistics that estimate them, as in Tutorial 10's table.
:::

## Are Rural Counties More Likely to Support Trump?

Tutorial 7 took up the familiar claim that rural voters delivered Trump's 2016 victory, and found what it would predict: Trump's share of counties climbed steadily with rurality, and urban and rural counties differed by nearly 20 percentage points. The theoretical explanation was his campaign's emphasis on lost manufacturing jobs, the trade agreements blamed for those losses, and regulation of coal --- appeals with the most purchase where those industries anchor the local economy. But a cross-tab describes what these counties did in one election. Does that gap reflect a relationship between rurality and how counties vote, or is it the kind of gap that could appear even without one?

The hypothesis, unchanged from Tutorial 7:

::: {.hypothesis}
In a comparison of [counties], those that are [more rural] will be more likely to [give a majority of their support to Trump] than will those that are [less rural (more urban)].
:::

The data frame **counties** contains **TrumpMajority**, stored as a numeric variable taking the value 0 if Clinton won a majority of the vote and 1 if Trump won a majority, and **rural**, a 3-category ordinal variable also stored as a number, where 0 is more urban, 1 is neither urban nor rural, and 2 is more rural. Neither variable has any missing values, so all 3,112 counties appear in the table.

The cross-tab you built in Tutorial 7 already contains everything needed for the chi-square test --- we just need to add two arguments to the call. We start with the general claim here, then pause to introduce the second tool before returning to test the directional one.

### Testing the General Claim

#### Step 1: State the Null Hypothesis

$H_0:$ **rural** and **TrumpMajority** are not related in the population.

#### Step 2: State the Alternative Hypothesis

$H_A:$ **rural** and **TrumpMajority** are related in the population.

#### Step 3: Run the Test

**Run and observe 1: The chi-square test of independence**

The [descr]{.package-name} package must be loaded with `library(descr)` before you can use `CrossTable()`. Once a package is loaded, all of its functions are available for the rest of your R session --- you do not need to load it again each time you use a function from it. The exercise below loads it once; you can assume it is already loaded for all subsequent exercises in this tutorial.

**TrumpMajority** is assigned to `x` because it is our outcome variable --- the one whose distribution we are comparing across groups --- and **rural** is assigned to `y` because it is our explanatory variable, defining the columns across which we read column proportions.

```{r crosstab-trumpmajority-by-rural, exercise=TRUE, exercise.cap="Run and observe 1"}
library(descr)
CrossTable(x = counties$TrumpMajority,
           y = counties$rural,
           prop.r = FALSE,
           prop.c = TRUE,
           prop.t = FALSE,
           prop.chisq = FALSE,
           chisq = TRUE,
           expected = TRUE)
```

Look at the output carefully before moving on. Each cell now shows three numbers stacked vertically: the observed count, the expected count under independence (the number that would appear if the two variables were completely unrelated), and the column proportion you already know how to read. Notice too that the columns are headed 0, 1, and 2 rather than Urban, Neither, and Rural --- `CrossTable()` labels them with the values stored in **rural**, so you need the coding in mind as you read: 0 is the urban column and 2 is the rural one. The chi-square statistic and its p-value appear at the bottom of the output, in the section labeled "Pearson's Chi-squared test." That is where you will look in Steps 4 and 5 to reach a decision.

**Practice 2: Interpret an expected count**

```{r quiz-rural-expectedcounts, echo=FALSE}
question("A student looks at the Clinton-majority/urban cell in your output, sees an observed count of 303 against an expected count of 183.0, and says: \"So we expected only 183 urban counties to vote for Clinton, and 120 more of them did than should have.\" What is wrong with that reading?",
  correct = "Right. An expected count is a benchmark, not a forecast.",
  answer("Nothing was predicted --- 183.0 is what that cell would hold if the two variables were unrelated",
         correct = TRUE,
         message = "The expected count is built from the row and column totals alone: 490 Clinton-majority counties, 1,162 urban counties, 3,112 in all, giving (490)(1162)/3112 = 183.0. Nothing about how rurality and vote choice actually go together enters that calculation --- that is the point of it. It describes a world in which the two variables have nothing to do with each other, so that the observed 303 has something to be measured against. The gap of 120 is not an error or a surplus. It is how far this cell departs from independence, and adding up those departures across all six cells is what the chi-square statistic does."),
  answer("Nothing is wrong --- 183.0 is the count the chi-square test predicts for that cell",
         message = "The chi-square test predicts nothing. It builds a comparison figure from the margins, on the assumption that the two variables are unrelated, and then measures how far the data sit from it. Calling that figure a prediction turns the benchmark into a claim about what the world is like, which is what the test is trying to evaluate rather than assert."),
  answer("The expected count is calculated from the relationship between the two variables, so comparing it to the observed count is circular",
         message = "It is calculated from the row and column totals only. Those totals say how many counties are urban and how many went for Clinton, but nothing about whether the two go together --- which is exactly why the comparison is not circular."),
  answer("No cell can hold a fraction of a county, so 183.0 cannot be compared to 303",
         message = "Expected counts are usually not whole numbers, and that is fine. 183.0 is not a claim that 183 particular counties exist somewhere; it is the arithmetic benchmark the observed count is measured against. Rounding it to 183 would make the chi-square statistic slightly wrong for no gain."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  try_again = "The student's reading treats 183.0 as something the analysis forecast. Ask instead where that number came from --- which totals went into it, and what was deliberately left out."
)
```

Now that you can see how far the observed counts stray from what independence would predict, let's look at what the test statistic and p-value tell us about whether that gap would be surprising if the two variables were truly unrelated.

#### Steps 4 and 5: Find the P-value, Decide, and Report the Result

**Apply 3: Interpret the chi-square result**

```{r quiz-rural-general-result, echo=FALSE}
question("Find the p-value in your output. Which statement best summarizes what this test tells us?",
  correct = "Right. A chi-square p-value carries no direction.",
  answer("We reject the null --- evidence that **rural** and **TrumpMajority** are related --- but that evidence carries no direction",
         correct = TRUE,
         message = "The p-value is well below .05, so a pattern this strong would be very unlikely if the two variables were unrelated in the population. That is what rejecting the null amounts to: not a demonstration that a relationship exists, but evidence against the claim that none does. And notice what the test statistic was built from --- squared gaps between observed and expected counts. Squaring throws the sign away, so a table where Trump support rises with rurality and a table where it falls produce the same chi-square. The evidence points to a relationship without pointing anywhere in particular, which is why a hypothesis predicting a direction needs a second, targeted test."),
  answer("We reject the null, which proves that more rural counties are significantly more likely to support Trump than more urban counties",
         message = "Two problems. A hypothesis test never proves anything --- it tells you how unlikely the data would be if the null were true, which is evidence, not proof. And a chi-square result carries no direction at all, so it cannot support a claim about which group is higher."),
  answer("We fail to reject the null, so there is no relationship between **rural** and **TrumpMajority**",
         message = "A p-value this small leads us to reject the null, not fail to reject it. Failing to reject would require a p-value above .05 --- and even then it would not show that no relationship exists, only that the data gave us no strong evidence against that possibility."),
  answer("We reject the null, and since the relationship runs across all three rurality categories, the directional claim is supported too",
         message = "Running across all three categories is not the same as running in a particular direction. The chi-square statistic is built from squared gaps, so it would be just as large if Trump support fell with rurality. The direction is visible in the column proportions --- but it is the proportions test, not this one, that weighs evidence for it."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  try_again = "Find the p-value and compare it to .05. Then ask what the chi-square statistic was built from, and whether that arithmetic could distinguish a rising pattern from a falling one."
)
```

The general claim is settled. Testing the directional claim requires a second tool, which the next section introduces.

## The Difference-of-Proportions Test

### The Logic of the Test

A difference-of-proportions test states a null that two population proportions are equal, and a directional alternative that one is greater or less than the other. The parameter is $\pi$ (pi), the population proportion --- the Greek-letter counterpart to the sample proportion $p$.

The chi-square test has answered the general question. The directional claim needs a different tool --- one that compares exactly two groups and can test which way the difference runs. A [**difference-of-proportions test**]{.important-text}, run with `prop.test()`, compares the proportion of "successes" in one group to the proportion of "successes" in another. It compares exactly two groups at a time. Unlike the chi-square test, it can test a specific direction --- a one-tailed test, when a hypothesis predicts which group should have the higher proportion. Because building a cross-tab already gives us the counts a chi-square test needs, our usual approach will be to test the general claim first --- are these two variables related at all --- and then, when the hypothesis makes a specific, directional prediction about two groups, sharpen to a difference-of-proportions test as well.

### Running the Test with `prop.test()`

Unlike `CrossTable()`, `prop.test()` does not work directly with the raw data --- instead, we supply the counts ourselves, using the cross-tab we already built.

`prop.test()` compares exactly two groups. Which two groups is a question the hypothesis answers, not the variable. A group may be a single category, or several categories added together, and a category the hypothesis names on neither side takes no part in the test. `x` and `n` are two parallel lists with one entry per group: `x` holds each group's number of successes and `n` holds each group's total, so a group's proportion is its entry in `x` over its entry in `n`. Group 1's count must appear first in *both* vectors; Group 2's count must appear second in *both*. If you mismatch the order in one vector but not the other, R will not give an error --- it will just silently compare the wrong numbers.

The `alternative` argument controls whether the test is one-tailed or two-tailed. The default, `alternative = "two.sided"`, tests whether the two proportions differ in either direction. When a hypothesis predicts a specific direction --- as every hypothesis in this tutorial does --- we instead use `alternative = "less"` or `alternative = "greater"`, matching the direction the hypothesis predicts. The key rule: `alternative = "less"` tests whether the *first* group's proportion is lower than the *second's*; `alternative = "greater"` tests whether the first group's proportion is higher. Which group you list first in `x` and `n` therefore determines which `alternative` setting matches your hypothesis.

::: {.caution}
**Position carries the meaning in `prop.test()`.** Two things depend on the order you list the groups in, and R checks neither.

First, `x` and `n` must name the groups in the same order. List one group first in `x` and the other first in `n`, and R pairs a success count with the wrong total, producing two proportions that describe neither group. Nothing about that is invalid arithmetic, so no error and no warning appears --- only a result that may point the wrong way.

Second, `alternative` is read against that same order. `"less"` and `"greater"` refer to whichever group is listed first, not to the group your hypothesis names first. Swap the groups without swapping `alternative` and you test the opposite of what you intended.
:::

Putting those together, a two-group `prop.test()` call takes three arguments: the two count vectors and the direction.

**Example code:**

```{r proptest-template, exercise=FALSE, eval=FALSE}
prop.test(x = c(successes_group1, successes_group2),
          n = c(total_group1, total_group2),
          alternative = "less")
```

## Rural Counties: Testing the Directional Claim

Our hypothesis does not just predict that **rural** and **TrumpMajority** are related --- it predicts a specific direction: Trump support should be higher among rural counties than among urban counties specifically. Looking at the cross-tab, we can see that direction in the column proportions. But the chi-square p-value does not directly test that specific one-tailed contrast. To do that, we compare the two endpoint categories --- urban and rural. The hypothesis contrasts the two ends of the rurality scale, so the 1,323 counties that are neither urban nor rural belong to neither group and take no part in this test. Nothing is being discarded from the analysis: the chi-square test above used all 3,112 counties to ask whether the two variables are related at all, and this narrower test uses the two groups the hypothesis actually names.

### Step 1: State the Null Hypothesis

$H_0: \pi_{urban}=\pi_{rural}$, where $\pi$ is the proportion of counties with a Trump majority in the population.

### Step 2: State the Alternative Hypothesis

$H_A: \pi_{urban} < \pi_{rural}$

### Step 3: Run the Test

Reading from the cross-tab you just produced above, Trump won a majority in 859 of the 1162 urban counties, and in 584 of the 627 rural counties. Those are the successes and the trials the mechanics section described: `x` takes the two Trump-majority counts, `n` the two column totals, with urban listed first in both.

**Run and observe 4: A one-tailed difference-of-proportions test**

`prop.test()` is part of base R --- it does not require loading any package. We set `alternative = "less"` because the hypothesis predicts that the urban proportion is *lower* than the rural proportion --- not just different from it.

```{r proptest-trumpmajority-urban-rural, exercise=TRUE, exercise.cap="Run and observe 4"}
prop.test(x = c(859, 584), n = c(1162, 627), alternative = "less")
```

The output labels the test statistic `X-squared`, which is the same label you saw on the chi-square test above. That is not because this is the same test. It answers a different question, on two groups rather than all three, and it is testing a direction the chi-square test cannot. The label is there because `prop.test()` calculates its p-value using a chi-square approximation.

The output also reports a 95% confidence interval --- here (-1.0000000, -0.1640333), for the difference in proportions, urban minus rural. One bound looks strange because this is a one-tailed test: R fixes it at the theoretical extreme rather than calculating it from the data, and for a difference in proportions that extreme is -1, since a difference cannot go below that. That is expected, not an error. The bound that matters is the calculated one, -0.164, and what to ask of it is whether the interval excludes 0. The fixed bound does not settle that by itself --- had the calculated bound come out at 0.05, the interval would read (-1, 0.05), would include 0, and the test would fail to reject the null. When a confidence interval and a hypothesis test are matched, meaning the same significance level and the same one- or two-tailed specification, they agree; the decision in this tutorial rests on the p-value, with the interval worth a glance as a cross-check.

### Steps 4 and 5: Find the P-value, Decide, and Report the Result

**Apply 5: Evaluate the rural/Trump-support hypothesis**

```{r quiz-rural-directional-result, echo=FALSE}
question("Find the p-value in your output. Does this result provide statistically significant evidence in support of the hypothesis?",
  correct = "Right. Significant, and in the direction the hypothesis predicted.",
  answer("Yes --- the p-value is far below .05 and the urban proportion is the lower one, matching the prediction",
         correct = TRUE,
         message = "Trump won a majority in 73.9% of urban counties against 93.1% of rural ones, and the p-value is below .001, so a gap this size would be very unlikely if the population proportions were equal. The confidence interval agrees, excluding 0. Taken with the chi-square test, this is statistically significant evidence in support of the hypothesis --- the chi-square gave evidence of some relationship, and this test gives evidence that the difference runs the way the hypothesis predicted. Neither result proves the hypothesis; together they are what support for it looks like."),
  answer("Yes, but only weakly --- the p-value is small because the sample is large, not because the effect is meaningful",
         message = "Sample size does affect the p-value, but that does not make this result meaningless. The proportions themselves --- 73.9% versus 93.1% --- represent a 19-percentage-point difference between urban and rural counties. Both the statistical significance and the magnitude of the difference are relevant to evaluating the hypothesis."),
  answer("We cannot conclude anything about direction from this test alone --- prop.test() only tells us whether the proportions differ, not which is higher",
         message = "prop.test() does tell you which proportion is higher --- the sample proportions appear in the output, and alternative = \"less\" specifically tests whether the first group's proportion is lower. The direction is right there in the output."),
  answer("No --- the p-value is below .05, so we fail to reject the null hypothesis",
         message = "This reverses the decision rule. A p-value *below* the significance threshold (.05) leads us to *reject* the null hypothesis, not fail to reject it. We fail to reject the null only when the p-value is *above* .05 --- meaning the observed pattern would not be surprising if the null were true."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  try_again = "Hint: Find the p-value and check whether it is above or below .05. Then check whether the sample proportions in the output match the direction the hypothesis predicted."
)
```

## Are Majority-White Counties More Likely to Support Trump?

Tutorial 7 found the sharpest contrast of its three examples here: White Majority counties gave Trump a majority far more often than counties with any other racial majority. The theoretical explanation was the realignment that followed the Civil Rights and Voting Rights Acts, which made race a durable dividing line in American elections and left it unusually visible in 2016. Does that contrast reflect a relationship between race and how counties vote, or is it the kind of contrast that could appear even without one?

The hypothesis, unchanged from Tutorial 7:

::: {.hypothesis}
In a comparison of [counties], those that are [majority-white] will be more likely to [give a majority of their support to Trump] than will those that are [not majority-white].
:::

The data frame **counties** contains **TrumpMajority**, again stored as a numeric variable coded 0 and 1, and **racial_majority**, a nominal variable stored as a character variable that takes the values "Black Majority," "Hispanic Majority," "Other/No Majority," and "White Majority." Neither variable has any missing values, so all 3,112 counties appear in the table.

As in the previous example, we start with the general claim that the cross-tab can answer for free, then sharpen to the directional claim the hypothesis actually makes.

### Testing the General Claim

#### Step 1: State the Null Hypothesis

$H_0:$ **racial_majority** and **TrumpMajority** are not related in the population.

#### Step 2: State the Alternative Hypothesis

$H_A:$ **racial_majority** and **TrumpMajority** are related in the population.

#### Step 3: Run the Test

The cross-tab you built in Tutorial 7 already contains everything the chi-square test needs --- just add `chisq = TRUE` and `expected = TRUE` to the call you already know.

**Practice 6: Add the chi-square test to the racial-majority cross-tab**

Fill in the two blanks with **TrumpMajority** and **racial_majority**, both from the **counties** data frame. The hypothesis tells you which of the two is the outcome, and the outcome goes in `x`. Two hints are available. Click "Submit Answer" to check your work.

```{r crosstab-trumpmajority-by-racialmajority, exercise=TRUE, exercise.cap="Practice 6"}
library(descr)
CrossTable(x = counties$___,
           y = counties$___,
           prop.r = FALSE,
           prop.c = TRUE,
           prop.t = FALSE,
           prop.chisq = FALSE,
           chisq = TRUE,
           expected = TRUE)
```

```{r crosstab-trumpmajority-by-racialmajority-hint-1}
# Hint 1 of 2
# Only two things are blank, and both are variable names. The table should
# show how counties divided between Clinton and Trump within each racial
# majority category --- so one variable is the outcome you are comparing,
# and the other defines the groups you compare it across. Read the
# hypothesis again: which of the two is doing the explaining?
```

```{r crosstab-trumpmajority-by-racialmajority-hint-2}
# Hint 2 of 2
# The hypothesis predicts that majority-white counties are more likely to give
# Trump a majority --- so TrumpMajority is the outcome and goes in x, and
# racial_majority is the explanatory variable and goes in y.
```

```{r crosstab-trumpmajority-by-racialmajority-check}
grade_this({
  code <- paste(.user_code, collapse = "\n")
  code_no_comments <- gsub("#[^\r\n]*", "", code)

  if (is.null(.result)) {
    fail("Your code did not produce output. Check for unmatched parentheses, a missing comma between arguments, or a misspelled argument name.")
  }
  if (!grepl("CrossTable\\s*\\(", code_no_comments)) {
    fail("Use `CrossTable()` to create the cross-tab.")
  }
  if (!grepl("TrumpMajority", .result$RowData)) {
    fail("The outcome variable goes in the `x` argument, so it appears in the rows. Which of the two does the hypothesis treat as the outcome?")
  }
  if (!grepl("racial_majority", .result$ColData)) {
    fail("The explanatory variable goes in the `y` argument, so it appears in the columns.")
  }
  if (!grepl("prop\\.r\\s*=\\s*FALSE", code_no_comments)) {
    fail("Set `prop.r = FALSE` to suppress row proportions.")
  }
  if (!grepl("prop\\.c\\s*=\\s*TRUE", code_no_comments)) {
    fail("Set `prop.c = TRUE` to show column proportions.")
  }
  if (!grepl("prop\\.t\\s*=\\s*FALSE", code_no_comments)) {
    fail("Set `prop.t = FALSE` to suppress total proportions.")
  }
  if (!grepl("prop\\.chisq\\s*=\\s*FALSE", code_no_comments)) {
    fail("Set `prop.chisq = FALSE` to suppress chi-square cell contributions.")
  }
  if (!isTRUE(.result$chisq)) {
    fail("Set `chisq = TRUE` to add the chi-square test of independence to the output.")
  }
  if (!isTRUE(.result$expected)) {
    fail("Set `expected = TRUE` to add the expected counts to the output.")
  }

  pass("Exactly! **TrumpMajority** is assigned to `x` and **racial_majority** is assigned to `y`, with the chi-square test and expected counts included.")
})
```

You will notice the output now has four columns instead of three, and each cell shows three numbers. The structure is the same as in the first example: observed count, expected count, column proportion. With four racial majority categories and two outcome categories, the table is larger, but the reading logic is identical.

#### Steps 4 and 5: Find the P-value, Decide, and Report the Result

The p-value appears at the bottom of the `CrossTable()` output, in the Pearson's Chi-squared section. Find it, compare it to .05, and then decide what this test --- which treats all four racial majority categories at once --- can and cannot tell you.

**Apply 7: Diagnose an overstated conclusion**

```{r quiz-racialmajority-general-result, echo=FALSE}
question("A student reads the chi-square output and says: 'p < .001, so we have strong evidence that majority-white counties are more likely to support Trump.' What is wrong with this conclusion?",
  correct = "Right. A significant chi-square does not say which categories drove it.",
  answer("The test pools all four racial-majority categories, so it cannot say which of them the evidence is about",
         correct = TRUE,
         message = "This table has four columns, and the chi-square statistic adds up departures from independence across every cell of every one of them. A large total can arise many ways. It could come from White Majority counties standing apart, as the student assumes --- or from Black Majority counties being unusually Clinton-leaning, or from Hispanic and Other/No Majority counties differing from each other, or from any mixture. The single p-value cannot separate these, because every cell contributed to it. Reading the column proportions will show you where the action is; the proportions test is what weighs evidence for a claim about one specific comparison."),
  answer("Nothing is wrong --- the column proportions show White Majority counties highest, and the test is significant",
         message = "The proportions do show that, and they are worth reading. But they describe this sample, while the test is what asks whether a pattern like it would be unlikely in the population --- and the test asked about all four categories at once. Setting a descriptive reading beside a significant omnibus p-value does not produce tested evidence for a claim about one category. Getting that requires a test aimed at the comparison you want to make, which is what the proportions test does."),
  answer("The conclusion is wrong because the chi-square test has not yet been confirmed by a proportions test",
         message = "The chi-square result is sound for the question it asked. The trouble is not that it awaits confirmation but that it answered a different question --- one about all four categories together rather than about White Majority counties in particular."),
  answer("The conclusion is too narrow --- a significant chi-square shows that all four categories differ from one another",
         message = "Rejecting this null says only that the four categories are not all alike. It does not say that each differs from every other. A significant result is consistent with three categories being nearly identical while one stands apart, and with a good many other arrangements besides --- it identifies neither how many categories differ nor which ones."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  try_again = "The student\'s claim names one category. Count how many the test used, and ask whether one p-value built from all of them can be traced back to any one."
)
```

### Testing the Directional Claim

Our hypothesis is specifically about majority-white counties compared to *all other* counties combined --- not about comparing White Majority counties to just one other category, and not about the four categories generally. This comparison group does not exist as a single column in the cross-tab --- we have to build it by adding across the three non-white-majority categories. Where the rural hypothesis named two poles and left the middle category out, this one names a group spanning three categories, so here we build the group rather than setting anything aside.

`prop.test()` needs four numbers: the Trump-majority count and the total for each of the two groups. Both come from your cross-tab, but from different rows:

- **x** takes the Trump-majority counts --- these come from the row where **TrumpMajority** = 1. For White Majority counties, read that cell directly: 2549. For the combined non-white-majority group, find the **TrumpMajority** = 1 cell for each of the three non-white columns and add them: 4 + 55 + 14.
- **n** takes the column totals --- these come from the Total row at the bottom. For White Majority counties, read the column total directly: 2866. For the combined group, find the Total cell for each of the three non-white columns and add them: 102 + 96 + 48.

Rather than computing the sums in your head, write the addition directly inside the `prop.test()` call --- R evaluates arithmetic expressions anywhere you can pass a number, so `4 + 55 + 14` inside `c()` works exactly like typing `73`. The calculation stays visible and checkable in your code. Pass both pairs of numbers as vectors using `c()`: `x = c(White_count, combined_count)` and `n = c(White_total, combined_total)`, with the White Majority group listed first in both.

#### Steps 1 and 2: State the Hypotheses

$H_0: \pi_{White}=\pi_{NotWhite}$, where $\pi$ is the proportion of counties with a Trump majority in the population.

$H_A: \pi_{White} > \pi_{NotWhite}$

Before building the test, identify the correct `alternative` argument.

**Practice 8: Choose the correct alternative**

```{r quiz-racialmajority-alternative, echo=FALSE}
question("White Majority counties are listed first in `x` and `n`. Which value of `alternative` correctly tests the directional claim?",
  correct = "Right. First group listed, higher proportion predicted, so `greater`.",
  answer('`alternative = "greater"`',
         correct = TRUE,
         message = "`alternative` is read against the order you listed the groups in, not against the order the hypothesis names them. White Majority counties occupy the first position in both `x` and `n`, and the hypothesis predicts that group has the higher proportion --- so the question to put to R is whether the first proportion exceeds the second, which is `greater`. Had you built the vectors the other way round, the same hypothesis would need `less`. The setting follows the arrangement of your own code."),
  answer('`alternative = "less"`',
         message = 'This would test whether the White Majority proportion is lower than the combined-other proportion --- the opposite of what the hypothesis predicts. Check the direction again.'),
  answer('`alternative = "two.sided"`',
         message = 'A two-sided alternative tests whether the proportions differ in either direction. The hypothesis makes a specific directional prediction, so a one-tailed test is appropriate here.'),
  answer('`alternative = "greater"` would only be correct if the groups were listed in the opposite order',
         message = 'The choice of `alternative` depends on both the direction the hypothesis predicts and which group is listed first. With White Majority first and a hypothesis predicting it has the higher proportion, `alternative = "greater"` is correct as written --- it would need to change to `"less"` only if the groups were swapped.'),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  try_again = "Hint: The hypothesis predicts White Majority counties have a higher Trump-majority proportion. With that group listed first in x and n, does alternative = \"greater\" or \"less\" test that direction?"
)
```

#### Step 3: Run the Test

**Practice 9: Build the `prop.test()` call for racial majority**

Write the `prop.test()` call, with White Majority listed first. Two hints are available. Click "Submit Answer" to check your work.

```{r proptest-trumpmajority-white-other, exercise=TRUE, exercise.cap="Practice 9"}

```

```{r proptest-trumpmajority-white-other-hint-1}
# Hint 1 of 2
# Two groups, two proportions. For each group you need the number of counties
# that gave Trump a majority and the total number of counties in that group.
# White Majority counties give you both directly. The other group does not
# exist as a single column --- you build it by adding the three non-white
# columns together. One row of the cross-tab holds the Trump-majority counts
# and another holds the totals; find both before you write anything.
```

```{r proptest-trumpmajority-white-other-hint-2}
# Hint 2 of 2
# The Trump-majority counts come from the row where TrumpMajority = 1; the
# group totals come from the Total row at the bottom. List the White Majority
# pair first in both vectors. Then the direction: the hypothesis predicts
# White Majority has the higher proportion, and with that group listed first,
# which value of alternative tests it?
```

```{r proptest-trumpmajority-white-other-check}
grade_this({
  code <- paste(.user_code, collapse = "\n")
  code_no_comments <- gsub("#[^\r\n]*", "", code)

  if (is.null(.result)) {
    fail("Your code did not produce a result. Check for unmatched parentheses, values in the wrong order, or a missing argument.")
  }
  if (!inherits(.result, "htest")) {
    fail("Your code should return the `prop.test()` output. Check for unmatched parentheses, values in the wrong order, or a missing argument.")
  }
  if (!grepl("prop\\.test\\s*\\(", code_no_comments)) {
    fail("Use `prop.test()` to run the difference-of-proportions test.")
  }
  if (!grepl("alternative\\s*=\\s*['\"]greater['\"]", code_no_comments)) {
    fail('The hypothesis predicts the White Majority proportion is higher than the combined-other proportion. Set `alternative = "greater"` to test that specific direction.')
  }

  expected <- prop.test(x = c(2549, 73),
                        n = c(2866, 246),
                        alternative = "greater")

  if (!isTRUE(all.equal(as.numeric(.result$p.value),
                        as.numeric(expected$p.value),
                        tolerance = 1e-6))) {
    fail("Check your counts. `x` should give `c(2549, 73)` --- Trump-majority counties among White Majority, then among the other three categories combined (4 + 55 + 14). `n` should give `c(2866, 246)` --- totals for those same two groups (102 + 96 + 48 for the combined group).")
  }

  pass("Well done! White Majority counties are compared to all other counties combined, with `alternative = \"greater\"` matching the direction the hypothesis predicts. Writing the additions inside c(), rather than working the sums out first, keeps the calculation visible in your code.")
})
```

#### Steps 4 and 5: Find the P-value, Decide, and Report the Result

**Apply 10: Evaluate the racial-majority/Trump-support hypothesis**

```{r quiz-racialmajority-directional-result, echo=FALSE}
question("Find the p-value in your output. Does this result provide statistically significant evidence in support of the hypothesis?",
  correct = "Right. Significant, and in the predicted direction.",
  answer("Yes --- the p-value is far below .05 and White Majority counties have the higher proportion, as predicted",
         correct = TRUE,
         message = "Trump won a majority in 88.9% of White Majority counties against 29.7% of all other counties combined, and the p-value is below .001 --- a gap that size would be very unlikely if the population proportions were equal. The confidence interval agrees, excluding 0. This is where the chi-square test\'s limitation gets repaired: that test gave evidence of some relationship among the four categories without saying which, and this one weighs evidence for the particular comparison the hypothesis names."),
  answer("Yes, but the result may be driven by the very small number of non-white-majority counties --- with only 246 counties in that group, the p-value could be unreliable",
         message = "246 counties is not a particularly small comparison group for this test, and both the Trump-majority and non-Trump-majority counts within that group are large enough for the large-sample approximation to be reasonable. The difference is both statistically significant and substantively large."),
  answer("No --- the p-value is below .05, so we fail to reject the null hypothesis",
         message = "This reverses the decision rule. A p-value *below* .05 leads us to *reject* the null hypothesis, not fail to reject it. We fail to reject the null only when the p-value is *above* .05."),
  answer("We cannot tell from the output which direction the difference runs --- prop.test() only reports whether the proportions differ",
         message = "The sample proportions appear in the output. The direction is visible there, and alternative = \"greater\" specifically tested whether the White Majority proportion was higher."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  try_again = "Hint: Find the p-value. Then check whether the sample proportions in the output match the direction the hypothesis predicted: which group has the higher Trump-majority proportion?"
)
```

## Are Dictatorships Less Likely to Have High Gender Equality?

Tutorial 7 found dictatorships lagging behind both kinds of democracy on high gender equality. The theoretical explanation was that democratic institutions create pressure for equal rights, representation and policy responsiveness, while dictatorships face far less of it. But that comparison rested on only 61 countries with both measures recorded. Is the difference there in the underlying process, or could so few cases produce it by chance?

The hypothesis, unchanged from Tutorial 7:

::: {.hypothesis}
In a comparison of [countries], those that are [dictatorships] will be less likely to [have high gender equality] than will those that are [democracies].
:::

The data frame **world** contains **gender_equal3**, a 3-category factor whose levels are recorded as "Low," "Medium," and "High," and **regime_type3**, a 3-category factor whose levels are recorded as "Dictatorship," "Parliamentary democ," and "Presidential democ." Those abbreviated labels are what will appear in your cross-tab output. Both variables have substantial missing values: **regime_type3** is missing for 34 of the 167 countries and **gender_equal3** for 95. Only 61 countries have a recorded value for both, so that is the number your cross-tab will describe --- as in Tutorial 7, the 106 countries missing a value on either variable simply do not appear, with nothing in the output to mark their absence but the total.

That small number introduces a practical complication which comes up often in political science: with few cases spread across many cells, some expected counts fall below 5 and the standard chi-square approximation becomes unreliable. This example shows what to do when that happens.

As in the previous two examples, we start with the general claim the cross-tab can answer for free, then sharpen to the directional claim the hypothesis actually makes.

### Testing the General Claim

#### Step 1: State the Null Hypothesis

$H_0:$ **regime_type3** and **gender_equal3** are not related in the population.

#### Step 2: State the Alternative Hypothesis

$H_A:$ **regime_type3** and **gender_equal3** are related in the population.

#### Step 3: Run the Test

The structure is identical to the previous two examples --- only the variable names and data frame change. This time you will build the full `CrossTable()` call from scratch.

**Practice 11: Build the chi-square test for gender equality and regime type**

Build the full call using **gender_equal3** and **regime_type3**, both from the **world** data frame, and include `chisq = TRUE` and `expected = TRUE`. As before, the hypothesis tells you which variable is the outcome. Two hints are available. Click "Submit Answer" to check your work.

```{r crosstab-genderequality-by-regimetype, exercise=TRUE, exercise.cap="Practice 11", warning=TRUE}

```

```{r crosstab-genderequality-by-regimetype-hint-1}
# Hint 1 of 2
# You want a table showing, within each regime type, how countries divide
# across the three levels of gender equality --- together with the counts you
# would expect if regime type and gender equality had nothing to do with each
# other. The hypothesis predicts something about dictatorships, which tells
# you which of the two variables defines the groups being compared.
```

```{r crosstab-genderequality-by-regimetype-hint-2}
# Hint 2 of 2
# x = world$outcome.variable, y = world$explanatory.variable.
# gender_equal3 is the outcome; regime_type3 is the explanatory variable.
```

```{r crosstab-genderequality-by-regimetype-check}
grade_this({
  code <- paste(.user_code, collapse = "\n")
  code_no_comments <- gsub("#[^\r\n]*", "", code)

  if (is.null(.result)) {
    fail("Your code did not produce output. Check for unmatched parentheses, a missing comma between arguments, or a misspelled argument name.")
  }
  if (!grepl("CrossTable\\s*\\(", code_no_comments)) {
    fail("Use `CrossTable()` to create the cross-tab.")
  }
  if (!grepl("gender_equal3", .result$RowData)) {
    fail("The outcome variable goes in the `x` argument, so it appears in the rows. Which of the two does the hypothesis treat as the outcome?")
  }
  if (!grepl("regime_type3", .result$ColData)) {
    fail("The explanatory variable goes in the `y` argument, so it appears in the columns.")
  }
  if (!grepl("prop\\.r\\s*=\\s*FALSE", code_no_comments)) {
    fail("Set `prop.r = FALSE` to suppress row proportions.")
  }
  if (!grepl("prop\\.c\\s*=\\s*TRUE", code_no_comments)) {
    fail("Set `prop.c = TRUE` to show column proportions.")
  }
  if (!grepl("prop\\.t\\s*=\\s*FALSE", code_no_comments)) {
    fail("Set `prop.t = FALSE` to suppress total proportions.")
  }
  if (!grepl("prop\\.chisq\\s*=\\s*FALSE", code_no_comments)) {
    fail("Set `prop.chisq = FALSE` to suppress chi-square cell contributions.")
  }
  if (!isTRUE(.result$chisq)) {
    fail("Set `chisq = TRUE` to add the chi-square test of independence to the output.")
  }
  if (!isTRUE(.result$expected)) {
    fail("Set `expected = TRUE` to add the expected counts to the output.")
  }

  pass("Good work! **gender_equal3** is assigned to `x` and **regime_type3** is assigned to `y`, with `chisq = TRUE` and `expected = TRUE` added to the output.")
})
```

You should see a warning: "Chi-squared approximation may be incorrect." Two cells fall below the conventional threshold of 5 for expected counts --- **Dictatorship/Low** and **Presidential democ/Low**, both with an expected count of 4.7. This happens because both the Low gender equality row (only 18 countries) and the Dictatorship and Presidential democ columns (16 countries each) have small totals: (18)(16)/61 = 4.7. When row and column totals are both small, the expected count in that cell will be small too. With only 61 countries spread across nine cells, this is hard to avoid.

This is not a sign of a bad hypothesis or bad data --- it is just a situation where the standard tool needs a small adjustment. The fix is to ask R to calculate the p-value by simulation rather than by approximation, using `simulate.p.value = TRUE`. Because simulation draws random tables rather than reading a value off a fixed distribution, the p-value shifts a little every time the code runs. `set.seed()`, which you used in Tutorial 10, fixes the starting point of R's random number generator so that everyone running this code gets the same result.

**Run and observe 12: The test with a simulated p-value**

```{r crosstab-genderequality-by-regimetype-simulated, exercise=TRUE, exercise.cap="Run and observe 12"}
set.seed(2026)
CrossTable(x = world$gender_equal3,
           y = world$regime_type3,
           prop.r = FALSE,
           prop.c = TRUE,
           prop.t = FALSE,
           prop.chisq = FALSE,
           chisq = TRUE,
           expected = TRUE,
           simulate.p.value = TRUE)
```

Notice that this version of the output does not report degrees of freedom at all --- simulation does not rely on the chi-square distribution the way the standard approximation does, so there is no df to report.

#### Steps 4 and 5: Find the P-value, Decide, and Report the Result

The p-value is small enough to reject the null hypothesis, providing evidence that **regime_type3** and **gender_equal3** are related in the underlying process. As with the previous two examples, this chi-square result provides evidence of a relationship across all three categories of each variable --- it does not yet tell us whether dictatorships specifically lag behind democracies on high gender equality. We test that directional claim next.

### Testing the Directional Claim

Our hypothesis compares dictatorships to democracies generally --- not to one specific type of democracy. Just as we combined non-white-majority counties into a single group in the previous example, we combine both democratic regime types here to build a group that matches what the hypothesis actually predicts.

#### Steps 1 and 2: State the Hypotheses

This time, identify the complete null and alternative hypothesis pair before building the test.

**Practice 13: State the hypotheses**

```{r quiz-genderequality-hypotheses, echo=FALSE}
question("Which pair of hypotheses correctly formalizes the directional claim, with dictatorships listed first?",
  correct = "Right. Equal proportions in the null, the predicted direction in the alternative.",
  answer("$H_0: \\pi_{Dictatorship} = \\pi_{Democracy}$ and $H_A: \\pi_{Dictatorship} < \\pi_{Democracy}$",
         correct = TRUE,
         message = "The null is the no-difference baseline --- the two population proportions are equal --- and the alternative carries the prediction, that the proportion of dictatorships with high gender equality is the lower of the two. Both are written with $\\pi$, the population proportion, because both are claims about the population rather than about the 61 countries in the table. That is the pattern every test in this tutorial follows: the null says nothing is going on, the alternative says what you expect to find, and the test asks how unlikely your data would be if the null were the truth."),
  answer("$H_0: p_{Dictatorship} = p_{Democracy}$ and $H_A: p_{Dictatorship} < p_{Democracy}$",
         message = "The direction is right, but hypothesis tests are always stated in terms of population parameters, denoted with Greek letters. $\\pi$ (pi) is the population proportion --- the quantity the hypothesis is actually about. $p$ is the sample proportion that estimates it. The null and alternative describe what we believe about the population, not what we observe in the sample."),
  answer("$H_0: \\pi_{Dictatorship} = \\pi_{Democracy}$ and $H_A: \\pi_{Dictatorship} \\neq \\pi_{Democracy}$",
         message = "The null is right, but a two-sided alternative does not capture the directional prediction the hypothesis makes. The hypothesis specifically predicts dictatorships lag behind --- that is a one-tailed claim."),
  answer("$H_0: \\pi_{Dictatorship} < \\pi_{Democracy}$ and $H_A: \\pi_{Dictatorship} = \\pi_{Democracy}$",
         message = "In the hypothesis-testing framework used in this course, we state the null as the no-difference baseline --- equal population proportions --- and the alternative expresses the directional prediction we are testing. Swapping them here puts the directional claim in the null, which is not how the framework works."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  try_again = "Hint: What does the null hypothesis always say about the two population proportions? Then check which direction the hypothesis predicts --- and which letter denotes a population parameter."
)
```

#### Step 3: Run the Test

The chi-square test provided evidence of a relationship across the three regime types. Now we test the specific directional claim: that dictatorships have a lower proportion of countries with high gender equality than democracies. Because the hypothesis compares dictatorships to all democracies, the second group requires adding parliamentary and presidential democracies together --- the same cell-aggregation logic you used in the racial majority example.

**Practice 14: Build the difference-of-proportions test**

Read the counts for each group from the cross-tab you just produced, combining the two democratic categories. Two hints are available. Click "Submit Answer" to check your work.

```{r proptest-genderequality-dictatorship-democracy, exercise=TRUE, exercise.cap="Practice 14"}

```

```{r proptest-genderequality-dictatorship-democracy-hint-1}
# Hint 1 of 2
# The hypothesis compares dictatorships to ALL democracies combined ---
# both parliamentary and presidential. Find the high gender equality count
# and the total for each democratic category in your cross-tab, then add
# them together to build the second group. The first group is dictatorships alone.
```

```{r proptest-genderequality-dictatorship-democracy-hint-2}
# Hint 2 of 2
# The hypothesis predicts dictatorships have a LOWER proportion
# of high gender equality than democracies combined. With
# Dictatorship listed first in x and n, use the hypothesis
# to determine the correct alternative. For the counts, x takes
# the high gender equality count for each group and n takes the
# total --- use the values from your cross-tab and the addition
# described in Hint 1.
```

```{r proptest-genderequality-dictatorship-democracy-check}
grade_this({
  code <- paste(.user_code, collapse = "\n")
  code_no_comments <- gsub("#[^\r\n]*", "", code)

  if (is.null(.result)) {
    fail("Your code did not produce a result. Check for unmatched parentheses, values in the wrong order, or a missing argument.")
  }
  if (!inherits(.result, "htest")) {
    fail("Your code should return the `prop.test()` output. Check for unmatched parentheses, values in the wrong order, or a missing argument.")
  }
  if (!grepl("prop\\.test\\s*\\(", code_no_comments)) {
    fail("Use `prop.test()` to run the difference-of-proportions test.")
  }
  if (!grepl("alternative\\s*=\\s*['\"]less['\"]", code_no_comments)) {
    fail('The hypothesis predicts dictatorships have a lower proportion of high gender equality than democracies. Set `alternative = "less"` to test that specific direction.')
  }

  expected <- prop.test(x = c(1, 20),
                        n = c(16, 45),
                        alternative = "less")

  if (!isTRUE(all.equal(as.numeric(.result$p.value),
                        as.numeric(expected$p.value),
                        tolerance = 1e-6))) {
    fail("Check your counts against the cross-tab. `x` takes the High gender equality row: the Dictatorship cell, then the two democracy cells added together. `n` takes the Total row at the bottom for those same columns, again adding the two democracy totals.")
  }

  pass("Correct! Dictatorships are compared to all democracies combined, with `alternative = \"less\"` matching the direction the hypothesis predicts. The sample proportions are 6.3% for dictatorships and 44.4% for all democracies combined.")
})
```

#### Steps 4 and 5: Find the P-value, Decide, and Report the Result

**Apply 15: Evaluate the dictatorship/gender-equality hypothesis**

```{r quiz-genderequality-directional-result, echo=FALSE}
question("Find the p-value and the sample proportions in your output. Which statement best describes what this result tells us?",
  correct = "Right. Significant, in the predicted direction, and the gap is large.",
  answer("The p-value is below .05, dictatorships have the lower proportion as predicted, and the gap is substantively large",
         correct = TRUE,
         message = "Only 6.3% of dictatorships have high gender equality against 44.4% of democracies combined, with p = .007. Two separate things are worth saying about that. The p-value addresses whether a gap this size would be unlikely if the population proportions were equal, and at .007 it would be. The gap itself --- more than 38 percentage points --- addresses whether the difference matters, and a difference that size plainly does. A result can clear the first bar and not the second, so both are worth reporting; here they point the same way."),
  answer("The p-value is below .05 and the direction matches, but the gap between 6.3% and 44.4% is too small to be substantively meaningful",
         message = "A gap of over 38 percentage points between dictatorships and democracies on high gender equality is substantively large, not trivial. Statistical significance and substantive size point in the same direction here."),
  answer("No --- the p-value is below .05, so we fail to reject the null hypothesis",
         message = "This reverses the decision rule. A p-value *below* .05 leads us to *reject* the null hypothesis, not fail to reject it. We fail to reject the null only when the p-value is *above* .05."),
  answer("The p-value is below .05 but the direction is wrong --- dictatorships have a higher proportion of high gender equality than democracies",
         message = "Look at the sample proportions in your output: prop 1 (dictatorships) = 0.063 and prop 2 (democracies combined) = 0.444. Dictatorships have the lower proportion, which is exactly what the hypothesis predicted."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  try_again = "Hint: Find the p-value and compare it to .05. Then look at the sample proportions --- which group has the higher proportion of high gender equality, and does that match the predicted direction?"
)
```

## Testing a Single Proportion

The three examples above all tested whether two variables are related --- the general-then-directional procedure we used throughout. This section introduces the [**one-sample proportion test**]{.important-text}, a separate use of `prop.test()` that is not part of that procedure: it asks whether a single sample proportion differs from a fixed, hypothesized value, with no second variable and no relationship involved at all.

Its null names a value rather than a second group: $H_0: \pi = $ some specified proportion, tested against a two-sided or a directional alternative.

The cross-tab from the first example showed that Trump won a majority in most counties --- but "most" is imprecise. A more exact question: does Trump win a *majority of counties* in the statistical sense? That is, can we reject the null hypothesis that exactly half of all counties give Trump a majority? The hypothesized value is 0.50 --- not an arbitrary benchmark, but the threshold that separates "more counties than not" from "fewer counties than not." The research question is whether Trump's county-level dominance clears that bar by more than chance alone would produce.

We use the same `prop.test()` function, but the call has a different shape. In the two-group tests, `x` and `n` were each a vector built with `c()`, one entry per group. Here there is only one group, so each is a single number: `x` is the count of successes and `n` is the total. In place of a second group's counts we supply `p`, the hypothesized population proportion we are comparing against --- R's name for this argument, unrelated to the p-value despite sharing a letter. The numbers 2622 and 3112 come from the first example's cross-tab, and because the research question asks whether the proportion is *greater than* 0.50, we use `alternative = "greater"`.

**Example code:**

```{r proptest-onesample-template, exercise=FALSE, eval=FALSE}
prop.test(x = successes,
          n = total,
          p = hypothesized_proportion,
          alternative = "greater")
```

### Step 1: State the Null Hypothesis

$H_0: \pi = 0.50$, where $\pi$ is the population proportion of counties in which Trump won a majority.

### Step 2: State the Alternative Hypothesis

$H_A: \pi > 0.50$

### Step 3: Run the Test

**Run and observe 16: A one-sample proportion test**

```{r proptest-onesample-trumpmajority, exercise=TRUE, exercise.cap="Run and observe 16"}
prop.test(x = 2622, n = 3112, p = 0.50, alternative = "greater")
```

### Steps 4 and 5: Find the P-value, Decide, and Report the Result

The sample proportion is $p = 0.843$ --- Trump won a majority in 84.3% of counties. The test compares that proportion to the null value of 0.50. The p-value is below .001, so we reject the null hypothesis. The one-sided 95% confidence interval reported by R has a lower bound well above 0.50, pointing to the same conclusion: the population proportion of Trump-majority counties is significantly greater than half. Trump's county-level support clears the majority threshold by a wide and statistically significant margin.

## Check Your Understanding

The exercises below test whether you have internalized the key ideas from this tutorial: when cell-adding is needed, what a significant chi-square result can and cannot tell you, and how to translate a hypothesis into the correct `alternative` argument.

**Apply 17: Choose and set up the right test**

```{r cyu-recognize-celladding, echo=FALSE}
question("A hypothesis predicts that people with a college degree are more likely to support a policy than people without one, and education is measured with a 5-category variable (less than high school, high school, some college, bachelor's degree, graduate degree). Which approach correctly tests this hypothesis?",
  correct = "Right. Build each group from the categories the hypothesis puts in it.",
  answer("Combine bachelor's and graduate into one group and the other three into another, then run a one-tailed test with the degree group predicted higher",
         correct = TRUE,
         message = "Neither group the hypothesis names exists as a column of its own --- college degree spans two categories and no college degree spans three, so both sides have to be built by adding. This is the same move as the White Majority example, where the comparison group was assembled from three columns, except that here it happens on both sides at once. Then the direction: the hypothesis predicts the degree group is higher, so with that group listed first, `alternative = \"greater\"`. Which categories go where comes from the hypothesis; which `alternative` to use comes from the hypothesis together with the order you listed the groups in."),
  answer("Run a chi-square test of independence between education category and policy support",
         message = "A chi-square test would tell you whether education and policy support are related across all five categories, but its p-value does not directly test whether the college degree group specifically has a higher proportion of support than the non-college group. The hypothesis makes a specific directional comparison between two defined groups --- that requires a difference-of-proportions test."),
  answer("Run a one-tailed difference-of-proportions test comparing bachelor's degree holders to those with some college, since those are the two most comparable adjacent categories",
         message = "The hypothesis defines the comparison: college degree vs. no college degree. Choosing categories based on comparability rather than what the hypothesis predicts would change the question being asked. Both sides of the comparison require combining categories to match what the hypothesis actually predicts."),
  answer("Run a one-tailed difference-of-proportions test comparing the college degree group to the non-college group, with the non-college group predicted to have the higher proportion",
         message = "The direction is reversed. The hypothesis predicts that people with a college degree are more likely to support the policy --- so the college degree group should have the higher proportion, not the non-college group."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  try_again = "Hint: The hypothesis defines two groups that don't exist as single categories in the variable. It also predicts a specific direction. Both matter for choosing and setting up the right test."
)
```

**Apply 18: Distinguish a general result from a directional conclusion**

```{r cyu-general-vs-directional, echo=FALSE}
question("A researcher hypothesizes that military regimes are less likely than civilian regimes to hold competitive elections. A chi-square test comparing four regime categories and three election-competitiveness categories produces p < .001. Which conclusion is justified by that test alone?",
  correct = "Right. Evidence of some relationship, with neither a direction nor a culprit.",
  answer("That the two variables are related somehow --- the test names neither which regimes differ nor in which direction",
         correct = TRUE,
         message = "This is where both limits show at once. The test pooled four regime categories against three competitiveness categories, so the twelve cells all fed one statistic --- it cannot say the military category was the one that mattered. And the statistic squares every departure before adding, so it cannot say whether military regimes hold fewer competitive elections or more. What survives is that a pattern this strong would be very unlikely if the two variables were unrelated, which is evidence against that possibility rather than proof of any particular alternative. The researcher's hypothesis names one category and one direction, and needs a test that does the same."),
  answer("That military regimes have significantly less competitive elections than civilian regimes, since p < .001 is very strong evidence",
         message = "Strength of evidence is not scope of evidence. However small the p-value, the claim it bears on is the one the test made --- that no relationship holds anywhere across these categories --- and that claim names neither military regimes nor a direction."),
  answer("That the hypothesis is fully supported, since the test is highly statistically significant",
         message = "No test fully supports a hypothesis; a small p-value is evidence in support of one, never certainty. And this test bears on a broader claim than the hypothesis makes, so even taken as evidence it is evidence for something else."),
  answer("That the researcher made an error, since a chi-square test cannot be used when a hypothesis predicts a specific direction",
         message = "Running it was not an error. Starting with the general question is the sensible order, since a cross-tab supplies everything the chi-square needs --- the mistake would be stopping there and treating the answer as though it addressed the specific claim."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  try_again = "The hypothesis names one regime type and one direction. Ask whether a single p-value built from twelve cells of squared departures can speak to either."
)
```

**Apply 19: Predict the right direction for a new hypothesis**

```{r cyu-direction-prediction, echo=FALSE}
question('A hypothesis predicts that countries with lower levels of press freedom have a higher rate of high government corruption. You list the high-press-freedom group first in `x` and `n`. Which value of `alternative` tests the hypothesis directly?',
  correct = "Right. First group listed, lower proportion predicted, so `less`.",
  answer('`alternative = "less"`',
         correct = TRUE,
         message = "The hypothesis is stated about the low-press-freedom countries --- they have more corruption --- but the code is arranged the other way round, with the high-press-freedom group first. Restating the prediction to match that arrangement is the whole task: if low-press-freedom countries have more corruption, then high-press-freedom countries have less, so the first group's proportion is the lower one, and that is `less`. Notice that the hypothesis never changed. Only the sentence you had to translate it into did."),
  answer('`alternative = "greater"`',
         message = 'This would test whether the high-press-freedom group\'s proportion of high corruption is greater than the low-press-freedom group\'s --- the opposite of what the hypothesis predicts. The hypothesis says low-press-freedom countries have more corruption, which means high-press-freedom countries have less. With the high-press-freedom group listed first, `alternative = "less"` is the correct direction.'),
  answer('`alternative = "two.sided"`, because the test should account for uncertainty about which direction the relationship runs',
         message = "A two-sided test is appropriate when we have no directional prediction. The hypothesis makes a specific directional prediction --- lower press freedom goes with higher corruption --- so a one-tailed test is more appropriate. The only question is which direction to specify given which group is listed first."),
  answer('`alternative = "greater"`, because the low-press-freedom countries are the ones with more corruption, and they should go first',
         message = 'The group ordering is already fixed --- the high-press-freedom group is listed first. You cannot change that; you have to choose `alternative` to match the direction given that order. With high-press-freedom first and the hypothesis predicting it has lower corruption, `alternative = "less"` is correct.'),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  try_again = "Hint: The high-press-freedom group is listed first. The hypothesis predicts that group has lower corruption than the second group. Which value of alternative tests whether the first group's proportion is lower?"
)
```

## The Takeaways

The table below summarizes what a hypothesis test adds to the description you built in Tutorial 7; Tutorials 13 and 14 will fill in the remaining columns.

::: {.qpa-infographic}
<table class="qpa-table single">
  <caption>Testing Two Categorical Variables</caption>
  <thead>
    <tr>
      <th class="qpa-corner" scope="col"><span class="visually-hidden">Aspect of the test</span></th>
      <th scope="col">Two Categorical Variables</th>
    </tr>
  </thead>
  <tbody>
    <tr>
      <th scope="row">What the Test Asks</th>
      <td>Whether the variables are related at all; then whether two specific proportions differ</td>
    </tr>
    <tr>
      <th scope="row">Null Hypothesis</th>
      <td>The variables are independent; or &pi;<sub>1</sub> = &pi;<sub>2</sub></td>
    </tr>
    <tr>
      <th scope="row">Test Statistic</th>
      <td>Chi-square, for both tests</td>
    </tr>
    <tr>
      <th scope="row">R Functions</th>
      <td><code>CrossTable(chisq = TRUE)</code>, <code>prop.test()</code></td>
    </tr>
  </tbody>
</table>
:::

This tutorial turned the cross-tab comparisons from Tutorial 7 into formal **hypothesis tests**. Finding that column proportions differed in the direction a hypothesis predicted was a useful first look, but it could not tell us whether that difference was more extreme than we would expect if the null hypothesis were true. The two tests in this tutorial let us ask related but different questions: first whether the variables are related across all categories, and then whether a particular two-group contrast runs in the direction the hypothesis predicts.

This tutorial introduced three tests for categorical variables:

- The [**chi-square test of independence**]{.important-text} asks the general question: are these two categorical variables related at all in the population? Run it with `CrossTable()`, adding `chisq = TRUE` and `expected = TRUE` to the standard cross-tab call. The `expected = TRUE` argument adds the [**expected count**]{.important-text} --- (row total)(column total)/n --- to every cell, showing how far observed counts stray from what [**independence**]{.important-text} would predict, and the [**chi-square statistic**]{.important-text} sums those departures across the table. A chi-square test can provide evidence that a relationship exists, but its p-value does not directly test the specific directional contrast a hypothesis predicts. You can see the direction of the observed pattern in the cross-tab column proportions; the proportions test evaluates that particular directional claim directly. When R warns that the chi-square approximation may be incorrect because some expected counts are below 5, add `simulate.p.value = TRUE` to calculate the p-value by simulation instead of the standard approximation.

- The [**difference-of-proportions test**]{.important-text} asks the sharper, directional question: is the proportion of successes higher (or lower) in one specific group than in another? Run it with `prop.test(x = c(successes1, successes2), n = c(total1, total2), alternative = ...)`. The first entries in `x` and `n` must describe the same group; so must the second entries. Use `alternative = "less"` or `"greater"` when a hypothesis predicts a specific direction. When a hypothesis compares a group to all other categories combined, read the counts from the cross-tab and add across the cells you need before passing them to `prop.test()`.

- The [**one-sample proportion test**]{.important-text} asks a different question entirely: whether a single sample proportion differs from a fixed hypothesized value, with no second variable and no relationship involved. A natural use is testing whether a proportion clears a substantively meaningful threshold --- such as whether Trump won a majority of counties in the statistical sense, tested against a null of 0.50. Use the same `prop.test()` function with `p =` set to the hypothesized value instead of a second group's counts.

Our approach throughout was to begin with the general question and then sharpen the analysis to the specific directional comparison in the hypothesis. The chi-square test evaluates whether the variables are related across all categories; the proportions test directly evaluates the particular two-group directional contrast the hypothesis predicts.

The categorical variables in this tutorial were stored in three different ways: **TrumpMajority** as numeric (0 and 1), **racial_majority** as a character variable, and **gender_equal3** and **regime_type3** as factor variables. `CrossTable()` handles all three storage types without any special treatment. `prop.test()` works with counts you supply directly, so storage class does not affect it. Always check how a variable is stored before building a cross-tab or choosing an `alternative` --- factor variables sort by level order, character variables sort alphabetically, and numeric variables sort by value.

Every test follows the same [**five-step procedure**]{.important-text}: state $H_0$, state $H_A$, run the test, find the p-value, and decide whether to reject or fail to reject --- then report what that decision means substantively.

That decision is worth being precise about. Rejecting a null hypothesis --- even with a very small p-value --- never proves a hypothesis is true, and failing to reject it never proves it is false. A p-value tells us how surprising our result would be if the null were exactly correct; a small p-value is **statistically significant evidence in support of a hypothesis**, not certainty that it is right. That is not a small distinction. It is the most a hypothesis test can honestly tell us.

Tutorial 13 applies the same logic to a different combination of variable types: one categorical, one interval. The tool changes --- ANOVA for the general question, a t-test for the directional one --- but the procedure and the interpretive discipline stay exactly the same.

## Submit Your Work

When you have completed all exercises in this tutorial, generate your submission code using the button below. Copy the full code that appears and paste it into the Canvas submission box for this tutorial. The code can be long --- make sure you have copied all of it before submitting.

```{r encode-logic-12, context="server"}
learnrhash::encoder_logic()
```

```{r encode-ui-12, echo=FALSE}
learnrhash::encoder_ui(
  ui_before = shiny::div(
    shiny::p("Click the button below to generate your submission code.
              Copy the full code and paste it into Canvas.")
  )
)
```


```{r tutorial-success-banner-12, results="asis", echo=FALSE}
cat('
<p class="tutorial-complete"
   role="status"
   aria-label="Excellent work! You have completed this tutorial and can now test relationships between categorical variables using chi-square tests and difference-of-proportions tests.">
  <span aria-hidden="true">🎉</span>
  Excellent work! You can now test relationships between categorical variables using chi-square tests and difference-of-proportions tests.
</p>
')
```
